Oviductin sets the species-specificity of the mammalian zona pellucida

Author response:

The following is the authors’ response to the original reviews

Reviewer #3 (Recommendations for the authors):

Major concerns:

P.6, lines 223-224: The sentence sounds like the authors produced all the OVGP1s by themselves in their laboratories, which is not completely true. The recombinant human and mouse OVGP1s were purchased from OriGene. It is suggested that the authors should state and explain clearly here which OVGP1 is produced by their laboratories and that recombinant human and mouse OVGP1s were obtained and purchased from Origene.

It is already clearly included in the M&M.

P6, lines 227-229: The authors stated that "Western blots of the three OVGP1recombinants indicated expected sizes based on those of the proteins: 75 kDa for human and murine OVGP1 and around 60 kDa for bovine OVGP1 (Fig. 4B and S6)." I pointed out in my last review report that the size of the recombinant human OVGP1shown by the authors in their manuscript is not in agreement with what has been published previously in literature regarding the molecular weight of native human OVGP1 as well as that of recombinant human OVGP1. The authors did not address the above concern adequately. In fact, recombinant human OVGP1 has been produced a few years ago (Reproduction (2016) 152:561-573) and it has been previously demonstrated that a single protein band of approximately 110-130 kDa was detected for both native human OVGP1 (see Microscopy Research and Technique (1995) 32:57-69) and recombinant human OVGP1 (Reproduction (2016) 152:561-573; Carbohydrate Research (2012) 358:47-55) using antibodies specific for human OVGP1. Molecular weight of the protein core or polypeptide of human OVGP1 is approximately 75 kDa, but the glycosylated form of native human OVGP1 and recombinant human OVGP1 is approximately 110-130 kDa. Therefore, the authors might have been using the recombinant core protein of human OVGP1 instead of the fully glycosylated recombinant OVGP1 in their study. The same concern also applies to the commercially obtained mouse recombinant OVGP1 used by the authors in their study. I would also like to mention that the mature and fully glycosylated OVGP1s in mammals vary in molecular weight (90-95 kDa in domestic animals; 110-150 kDa in primates; 160-350 kDa in rodents). Again, the 75kDa of mouse OVGP1 detected by the authors could be the core protein or polypeptide of mouse OVGP1 instead of the fully glycosylated mouse OVGP1.

In our study, as previously mentioned, we included commercially available recombinant proteins from Origene for human and murine OVGP1, which are produced in mammalian cells, and we also produced and purified bovine OVGP1 in mammalian cells. Therefore, these proteins should be properly glycosylated. Moreover, we performed Western blot assays favouring the blotting of higher molecular weight proteins, ensuring the optimal conditions for the assay. Additionally, we tested the size of OVGP1 from murine and bovine oviductal fluids on the same blot. During oestrus, the size of OVGP1 from oviductal fluids matches that of the recombinant proteins, and this band is downregulated during anoestrus, confirming the proper size of recombinant protein.

P.7, lines 236 and 237: Please provide a figure or source to support the statement "...as confirmed by proteomics of the bands along with PEAKS Studio v11.5 search engine peptide identification software."

It is included in the text the amount of unique peptides obtained by Proteomics for OVGP1 identification over all protein groups identified.

P.7, lines 243 to 245: The statement "...using rabbit polyclonal antibody to human OVGP1 for bOVGP1 and endogenous OVGP1, and mouse monoclonal antibody against Flag (DDK)-tag for hOVGP1 and mOVGP1." is confusing and might be inaccurate. First of all, I wondered why the authors did not use an antibody against bovine OVGP1 for the recombinant bOVGP1 instead of using a rabbit polyclonal antibody to human OVGP1. Secondly, what does the "endogenous OVGP1" refer to in the statement? Thirdly, the authors in their study used the commercially available recombinant human OVGP1 and recombinant mouse OVGP1 purchased from Origene. Based on the data sheet provided by Origene, the tag used for both recombinant human OVGP1 and recombinant mouse is C-Myc/DDK-tag and not Flag-tag. Can the authors explain these discrepancies?

Firstly, for the recombinant protein of bOVGP1 we used the same antibody that we used in the Western blot for all the proteins and oviductal fluids because we do not have anti-His tag working for Immunofluorescence (the one we had only worked for Western blot) and neither we do not have any antibody against bovine OVGP1. In the case of human and murine since we had anti-Flag antibody that worked for Western blot and for immunofluorescence, we used this one. However, as has been shown in our figure and supplementary material, the antibody against human OVGP1 works properly for both techniques (Western blot and Immunofluorescence). Secondly, endogenous OVGP1 is referred to the OVGP1 present in the oviductal fluid. Thirdly, as you can see in the datasheet of the protein, the recombinant proteins purchased from Origene contains a c-myc tag (EQKLISEEDL) some amino acids and a ddk-tag (DYKDDDDK). The sequence of ddk is the same of Flag-tag (DYKDDDDK). Since the proteins have both tags we used the antibody against Flag (or ddk) epitope.

P12, lines 429-432: The newly added statement at the end of the Discussion saying "Additionally, future studies would be valuable to investigate whether incubating oocytes with oviductal fluid (or OVGP1) could reduce polyspermy in porcine IVF and whether ZPs could be leveraged to naturally enhance sperm selection in human ICSI" is very concerning and requires further attention. The statement reflects that the authors do not keep pace with and do not pay attention to what has been published in literature regarding porcine and human OVGP1s. In fact, porcine oviduct-specific glycoprotein (OVGP1) has already been reported to reduce the incidence of polyspermy in pig oocytes (Biology of Reproduction (2000) 63:242-250). Porcine oviductal fluid, used in porcine IVF, has also been found to exert a beneficial effect on oocytes by reducing the incidence of polyspermy without decreasing the penetration rate. (Theriogenology (2016) 86:495-502). Therefore, the studies deemed valuable by the authors to be investigated in the future have, in fact, already been carried out two decades ago by several other laboratories. I am surprised the authors were not aware of these published work in literature. All the above should have been incorporated in the Discussion.

This sentence is modified in the discussion and the references are included.

Furthermore, as mentioned earlier, recombinant human OVGP1 has also been produced (Reproduction (2016) 152:561-573), and recombinant human OVGP1 has been found to increase tyrosine phosphorylation of sperm proteins, a biochemical hallmark of sperm capacitation, and potentiate the subsequent acrosome reaction (Reproduction (2016) 152:561-573) as well as increase sperm-zona binding (Journal of Assisted Reproduction and Genetics (2019) 36:1363-1377). These earlier findings should be incorporated into the Discussion.

Thank you for your comment, but in this work we had not performed any experimental setting related to tyrosine phosphorylation and despite is a very interesting topic is not directly related to this work.

P.19, lines 678-683: Since the human and mouse recombinant oviductin proteins were purchased from Origene, the authors should be aware of the fact that these commercially available recombinant OVGP1s might not be fully glycosylated. While I appreciate the fact that the authors wanted to briefly describe how the human and mouse recombinant OVGP1s were prepared by the manufacturer, I strongly suggest that the authors should contact Origene, the manufacturer, for all information regarding the procedures for producing the human and mouse recombinant oviductin proteins. For example, the authors stated on lines 680-681 that "A sequence expressing FLAG-tagged epitope proteins (DYKDDDDK) was cloned into an expression vector." According to the data sheet provided by Origene, it appears that both human and recombinant oviductin proteins are C-Myc/DDK-tagged and not FLAG-tagged.

Thank you for your comment, as according to the sequence of Flag-tag it is matching with the sequence of the tag in the datasheet corresponding to DDK (this is in detail in previous comment). Besides, the protein is tagged also by C-Myc tag. Among both tags, the antibody selected to detect it was anti-Flag tag.

P.19, lines 692-697: The description of the primary and secondary antibodies used for detection of the various recombinant OVGP1s is also very confusing and not clearly presented. For example, it is mentioned here that "...membranes were...incubated with anti-OVGP1 rabbit monoclonal antibody for OVGP1,..". What specifically does "OVGP1" refer to here? The authors then stated that anti-Histamine Tag antibody was used to detect bOVGP1 and mOVGP1 and anti-Flag antibody was used to detect hOVGP1. As pointed out earlier, the human and mouse recombinant OVGP1s were produced using C-Myc/DDK tag and not His-tag or Flag-tag. Can the authors clarify these discrepancies?

We apologise for the complexity of the antibodies, we included in this paragraph the ones used to Western blot for both figures: anti- human OVGP1 was used for the principal figure that contains the three recombinant proteins and oviductal fluids; and the anti-Histidine and anti-Flag antibodies that are included in supplementary figure, specifically for recombinant bovine OVGP1 (Histidine tag) and for recombinant murine and human OVGP (DDK tag). A clarifying sentence has been included in the text.

P.31, lines 1143-1149: Figure 10 is not mentioned anywhere in the main text of the manuscript. Rewrite the second half of the sentence "...; being this specificity lost when OVGP1 is heterologous to the ZP (right diagram)." Which sounds awkward and grammatically not correct.

The figure is already mentioned in the text, thank you for your comment. The sentence is also corrected.

Other comments: P.1, the statement of "All authors contributed equally to this work" on line 14 can be deleted because detailed and specific contributions from each authors are listed in lines 1009-1017 on page 27.

Both authors contributed equally to this work, now is clear in authors contribution section.

P.2, lines 43 and 44: Do the authors mean "sperm-oocyte binding protein" instead of "sperm-oocyte fusion protein" in the sentence? "Fusion protein" is a protein composed of two or more domains encoded by different genes, or a hybrid molecule created by combining two different proteins for various purposes. I believe the term "fusion protein" is wrongly used in the sentence which should be rephrased with a proper term.

Done.

P2, line 73: Remove the comma after the word "Both".

Done.

P.5, line 179: "...mice ZP..." should be written as "...mouse ZP...".

Done.

P.6, heading of 3rd paragraph on line 207: The term "binding" will be a better term than "fusion" used in the heading because the results do not actually show the fusion of the OVGP1 proteins with the ZP glycoprotein. Instead, binding of the OVGP1 proteins to the ZP occurred.

Done.

P.6, lines 215-217: Authors, please provide a reference or references to support the statement "Region A, corresponding to the amino acid end, shows high identity among monotremes, marsupials and placentals."

In the text was indicated a review (29) which includes the supporting idea of this statement for Figure 4. Moreover, we have included some if the references used for the description of the domains when performing the sequence alignment of Figure S5.

P.6, line 230 and line 233 on P.7: Authors, please be consistent in the use of either American English or British English. The word "oestrus" is British English whereas "estrus" is American English.

Done.

P.7, line 264: The word "sticking" used here means non-specific binding. I believe the author means specific binding here. If so, a more appropriate word should be used here instead of "sticking".

Done.

P.7, lines 267-269: This newly added sentence sounds very awkward and should be completely rewritten.

Done.

P.8, line 288: This reviewer finds it difficult to understand the meaning of the heading. The heading should be rephrased to bring out exactly what the authors want to say in well-written English.

Done.

P.8, line 290: The word "would" should be replaced by "could" in the sentence.

Done.

P.13, line 437: Authors, please provide the location of Sigma-Aldrich.

Done.

P.13, line 457: Here, the authors used "1800 rpm" to indicate the centrifugation speed but used the g-force elsewhere in the Materials and Methods. Please be consistent. The g-force is preferred.

Done.

P.14, lines 483-485: The procedure of sacrificing the cats should be provided in the Materials and Methods

Cats weren’t sacrificed they were vasectomized. It is now included in the text.

P.17, line 628: "...the ZPs were exposed or no exposed to..." should be written as "...the ZPs were either exposed or not exposed to...".

Done.

P.17, line 629: "...each groups were incubated with..." should be "...each group was incubated with...".

Done.

P.19, line 700: "As loading control, was used the primary antibody....." is not a complete sentence and it needs to be rewritten.

Done.

P.20, lines 744-754: For scanning electron microscopy and image processing, the procedures of prior treatment of the oocytes with and without oviductal fluid and OVGP1 should be included here.

Done.

P.21, line 756: It is stated here that "Two hundred isolated ZPs were treated with Clostridium perfringens neuraminidase....". However, it is not clear whether two hundred isolated ZPs of both porcine and murine ZPs were treated. Authors, please clarify.

We used 200 isolated ZPs of each specie, bovine and murine. It is classified in the text.

P.28, lines 1039 and 1040: The author only mentioned the use of bovine and murine sperm here. What about human sperm?

Done.

P.29, line 1076: "...in mammalian cells..." is very vague. Be specific what exactly the mammalian cells were.

Done.

P.29, line 1079: "Oviductal fluid from ovulated cows or anoestrus cows." is not a complete sentence and it needs to be rewritten.

Done.

Changes from the earlier version to the revised version

Abstract

Abstract

The zona pellucida (ZP) is vital for species-specific fertilization as this barrier mediates sperm-oocyte binding. Here, we determined whether sperm from distant mammalian orders (Carnivora, Primates, and Rodentia) could penetrate bovine oocytes by examining the role of bovine oviductal fluid and species-specific oviductal glycoprotein (OVGP1 or oviductin) from bovine, murine, or human sources in modulating the species-specificity of bovine and murine oocytes. Sperm from all the species were found to penetrate intact bovine ovarian oocytes to form hybrid embryos. However, contact with oviductal fluid or bovine, murine, or human OVGP1, conferred the ZP species-specificity, allowing only the penetration of the corresponding sperm regardless of the ZP’s origin. Glycolytic and microstructural analyses revealed that OVGP1 covers the pores present in the ZP and that OVGP1 glycosylation determines sperm specificity. This suggests specific fertilization capacity is acquired in the oviduct through the ZP’s incorporation of specific oviductin.

Introduction

Earlier version: Interaction between the sperm and egg is generally assumed to be a species-specific event mediated by the recognition and binding of complementary molecules on the sperm plasma membrane to both the zona pellucida (ZP) and the oocyte oolemma (1, 2). However, it remains unclear whether both the oocyte’s ZP and plasma membrane are necessary to determine species-specificity. In evolutionary terms, the ZP is the oldest among the coatings that envelop vertebrate oocytes and conceptuses. This barrier is responsible for preventing polyspermy, facilitating preimplantation development, and mediating species-restricted recognition between the gametes (3). ZP proteins and sperm-oocyte fusion protein are usually found to evolve under selection pressure. It is proposed that this rapid evolution plays an important role in the reproductive isolation of diverging taxa (4, 5). The mechanism of fusion between the sperm and oolemma membranes is not fully understood, although various candidate molecules have been identified (6–8). The species-specificity of interactions between gamete membranes in mammals is not always apparent. The significance of the ZP in maintaining species boundaries is revealed by the fact that its removal allows for heterospecific sperm binding and penetration of the oocyte plasma membrane. For example, sperm from humans, mice, rats, guinea pigs, goats, bulls, horses, pigs, rabbits, and dolphins can all successfully penetrate hamster oocytes lacking their ZP coat (2, 9). This phenomenon was demonstrated in the hamster egg penetration test (HEPT) developed by Yanagimachi et al. (10). In this test, researchers observed that upon removal of the ZP from hamster ova, the eggs allowed for the penetration of sperm from other species (10). However, it remains unknown why zona-free hamster eggs are readily penetrated by heterologous spermatozoa, as this characteristic is not shared by the ova of other species. For instance, human spermatozoa will not penetrate zona-free rat, rabbit, or mouse ova (11). The ability of sperm from different species to penetrate the hamster oocyte suggests that the ZP could be a key factor in preventing heterologous fertilization between different species.

Revised version: Interaction between the sperm and egg is generally assumed to be a species-specific event mediated by the recognition and binding of complementary molecules on the sperm plasma membrane to both the zona pellucida (ZP) and the oocyte oolemma (1, 2). However, it remains unclear whether both the oocyte’s ZP and plasma membrane are necessary to determine species-specificity. In evolutionary terms, the ZP is the oldest among the coatings that envelop vertebrate oocytes and conceptuses. This barrier is responsible for preventing polyspermy, facilitating preimplantation development, and mediating species-restricted recognition between the gametes (3). ZP proteins and sperm-oocyte binding protein are usually found to evolve under selection pressure. It is proposed that this rapid evolution plays an important role in the reproductive isolation of diverging taxa (4, 5). The mechanism of fusion between the sperm and oolemma membranes is not fully understood, although various candidate molecules have been identified (6–8). The species-specificity of interactions between gamete membranes in mammals is not always apparent. The significance of the ZP in maintaining species boundaries is revealed by the fact that its removal allows for heterospecific sperm binding and penetration of the oocyte plasma membrane. For example, sperm from humans, mice, rats, guinea pigs, goats, bulls, horses, pigs, rabbits, and dolphins can all successfully penetrate hamster oocytes lacking their ZP coat (2, 9). This phenomenon was demonstrated in the hamster egg penetration test (HEPT) developed by Yanagimachi et al. (10). In this test, researchers observed that upon removal of the ZP from hamster ova, the eggs allowed for the penetration of sperm from other species (10). However, it remains unknown why zona-free hamster eggs are readily penetrated by heterologous spermatozoa, as this characteristic is not shared by the ova of other species. For instance, human spermatozoa will not penetrate zona-free rat, rabbit, or mouse ova (11). The ability of sperm from different species to penetrate the hamster oocyte suggests that the ZP could be a key factor in preventing heterologous fertilization between different species.

The successful in vitro fertilization (IVF) of zona-free bovine oocytes with deer spermatozoa has been reported (12). The IVF of zona-intact bovine oocytes with endangered African antelope spermatozoa was also described in 1998 (13). Further heterologous gamete interactions have been documented between porcine or equine sperm and ZP-intact bovine oocytes in large domestic animals (14), as have interspecific interactions between equine and porcine gametes (15). However, other research groups have reported complete species-specificity for IVF between heterologous gametes in pigs and cattle (16). In prior work, we observed in zona-intact bovine oocytes obtained directly from ovaries, the sperm fertility of other species of the same order, such as ram (17), common bottlenose dolphins (18), and wild goats (19). This suggests that species within the order Artiodactyla share similar fertilization mechanisms. However, it appears that this capacity is not restricted solely to this order. Both our group and other researchers have found that horse (order Perissodactyla) and guinea pig spermatozoa (order Rodentia) are also capable of fertilizing zona-intact bovine oocytes (20–22). In the present study, we used spermatozoa from the orders Carnivora (cats), Primates (humans), and Rodentia (mice, Muridae family) to examine whether species from other mammalian orders can also fertilize zona-intact bovine oocytes.

The ZP is an extracellular matrix that surrounds the oocyte and preimplantation embryo in mammals (23, 24). This matrix acts as a barrier, separating cumulus cells from the oocyte, facilitating species-restricted recognition during fertilization, and preventing polyspermy, contributing in this way to preimplantation embryo development (3). ZP matrix proteins are glycosylated and specific oligosaccharide residues have been linked to their role as sperm receptors (25). The diverse glycoconjugate composition of the ZP is thought to be unique to each species (3, 26). In cattle, the main components (77%) of the carbohydrate chains of ZP glycoproteins are acidic chains containing sialic acids (27): mainly Neu5Ac (84.5%) and Neu5Gc (15.5%) linked to Galβ1-4GlcNAc and GalNAc by α2,3-and α2,6-linkages on N-and O-glycans respectively (28). Oviductal fluid composition plays a crucial role in fertilization (29) and subsequent embryo development (30–32). Oviductal proteins have been proposed to induce changes in the ZP, influencing sperm-oocyte binding specificity (33, 34). These changes may impact ZP maturation, sperm-ZP binding, and the prevention of polyspermy (26). Modifications induced by oviductal proteins could alter the ZP’s composition by affecting its carbohydrates or proteins. One set of proteins found in oviductal fluid is a family of glycosylated proteins collectively known as oviduct-specific glycoprotein (OVGP1, also known as oviductin). This protein has been identified in various mammalian species, including humans. Oviductin has been associated with the ZP of ovarian oocytes post-ovulation (35). OVGP1, along with heparin-like glycosaminoglycans (GAGs), plays a role in functionally modifying the ZP in pigs and cows, making it more resilient to enzyme digestion and sperm penetration, thereby helping prevent polyspermy (36, 37). However, it remains uncertain whether the ZP’s sperm receptor is consistent across all mammalian species. This raises questions about how sperm from different mammalian orders are able to fertilize cow oocytes, and whether the lack of contact between cow oocytes from slaughterhouse ovaries and oviductal proteins prevents the changes in the ZP necessary for bull sperm recognition. Here, we have explored the role of bovine oviductal fluid and murine, human, or bovine oviductal glycoprotein 1 (OVGP1) in modulating the sperm-binding species-specificity of the bovine and murine ZP.

Results

Human, mouse, and cat spermatozoa can fertilize bovine ovarian oocytes

A schematic depiction of the experimental design is indicated in Figure 1.

Schematic representation of global experimental design of homologous and heterologous IVF from diverse mammalian species.

(A) Represents design of experiment shown in Tables S1-3, Figure 2A-B and Figure S1-3. (B) Represents design of experiment shown in Table S4, Figure 2C. (C) Represents design of experiment shown in Table S5-6, Figure 2D-E and Figure 3. (D) Represents design of experiment shown in Table S8-10, Figure 5 and Figure 7A. (E) Represents design of experiment shown in Table S11, and Figure 7B. (F) Represents design of experiment shown in Table S12-14, Figure 8-9.

We first examined whether human spermatozoa could penetrate bovine oocytes obtained directly from the ovary, without prior exposure to oviductal fluids. For this purpose, human spermatozoa were capacitated in two different media: G-IVF™ PLUS (HeA) and Fert (HeB), commonly used for the capacitation of human spermatozoa and bovine spermatozoa, respectively. As presented in Table S1 and depicted in Figures 2A and S1, human spermatozoa were able to penetrate the bovine oocytes, which went on to develop a pronucleus and eventually underwent cleavage into the two-cell stage. Heterologous fertilization outcomes were improved when the G-IVF™ PLUS medium (HeA) was used instead of Fert (HeB). However, these outcomes were still worse than those obtained for homologous fertilization. No differences were observed in the number of spermatozoa remaining attached to the ZP, or in pronuclear formation up to 12 hours post-incubation (hpi). These results indicate that human spermatozoa can bind to and penetrate both the bovine oocyte’s ZP and plasma membrane. Human spermatozoa were also found capable of inducing the first embryo cleavage.

Heterologous IVF of bovine oocytes, mouse oocytes, or empty mouse ZPs, using human, mouse, or cat sperm, before and after contact with oviductal fluid.

Earlier version: (A) Embryo cleavage rates resulting from the IVF of bovine oocytes with human, murine, or feline sperm including bovine sperm as homologous IVF control, and parthenogenesis as a negative control of the IVF. (B) Sperm penetration rate after the IVF of IVM mouse oocytes with murine or bovine sperm, for homologous and heterologous IVF, respectively. (C) Embryo cleavage rates resulting from the IVF of bovine IVM oocytes, preincubated 30 minutes with bovine oviductal fluid, with human sperm (IVF medium a = G-IVF™ PLUS medium (HeA); IVF medium b = Fert (HeB). (D) Penetration rates and average numbers of sperm bound to ZP after the empty zona penetration test (EZPT) using mouse ovarian IVM oocytes and murine or bovine sperm. (E) Penetration rates and average numbers of sperm bound to ZP after the EZPT using mouse oviductal oocytes and murine or bovine sperm. (F) Picture of an empty zona pellucida obtained from a mouse ovarian IVM oocyte after the EZPT using bovine sperm. Note the sperm has penetrated the zona. (G) Picture of an empty zona pellucida from a mouse oviductal oocyte after the EZPT using bovine sperm. Scale bar for ZP pictures = 50 µm. A non-fertilized parthenogenesis group is used as cleavage control in Fig. 2A and 2C. Different letters above error bars (mean ±SD) indicate significant differences (P < 0.05) among groups (ANOVA and Tukey’s post hoc test). Numbers of oocytes or ZPs used are indicated in Tables S1-S5).

Revised version: (A) Embryo cleavage rates resulting from the IVF of bovine oocytes with human, murine, or feline sperm including bovine sperm as homologous IVF control, and parthenogenesis as a negative control of the IVF. (B) Sperm penetration rate after the IVF of IVM mouse oocytes with murine or bovine sperm, for homologous and heterologous IVF, respectively. (C) Embryo cleavage rates resulting from the IVF of bovine IVM oocytes, preincubated 30 minutes with bovine oviductal fluid, with human sperm (IVF medium a = G-IVF™ PLUS medium (HeA); IVF medium b = Fert (HeB). (D) Penetration rates and average numbers of sperm bound to ZP after the empty zona penetration test (EZPT) using mouse ovarian IVM oocytes and murine, bovine or human sperm. (E) Penetration rates and average numbers of sperm bound to ZP after the EZPT using mouse oviductal oocytes and murine, bovine or human sperm. (F) Picture of an empty zona pellucida obtained from a mouse ovarian IVM oocyte after the EZPT using bovine sperm. Note the sperm has penetrated the zona. (G) Picture of an empty zona pellucida from a mouse oviductal oocyte after the EZPT using bovine sperm. Scale bar for ZP pictures = 50 µm. A non-fertilized parthenogenesis group is used as cleavage control in Fig. 2A and 2C. Different letters above error bars (mean ±SD) indicate significant differences (P < 0.05) among groups (ANOVA and Tukey’s post hoc test). Numbers of oocytes or ZPs used are indicated in Tables S1-S5).

As shown in Table S2 and illustrated in Figures 2A and S2, mouse spermatozoa showed a similar capacity to human spermatozoa to fertilize bovine ovarian oocytes. Following sperm penetration of the oocyte, the sperm nucleus decondensed, forming a pronucleus, which ultimately induced cleavage into two-cell stage bovine zygote. Pronucleus formation and cleavage rates were, however, lower than those obtained for homologous fertilization. Nonetheless, these findings confirm that mouse spermatozoa can also bind to and penetrate both the bovine oocyte ZP and plasma membrane, and are also capable of inducing the first embryo cleavage.

Similar results were also obtained with cat sperm. As presented in Table S3 and depicted in Figures 2A and S3, feline spermatozoa were found to bind to and penetrate bovine oocytes, thereby initiating the first cellular division of the hybrid embryo. Pronucleus formation and cleavage rates were lower than those observed for homologous bovine fertilization.

To explore if the same effect was produced using oocytes from another species, we collected oocytes from the ovaries of mice following PMSG-induced stimulation. The ovarian murine oocytes were matured in vitro and then fertilized with bovine sperm. The results obtained were similar to those obtained with bovine oocytes in that both the ZP and plasma membrane of the mouse oocytes could be penetrated by the bovine sperm, which then underwent decondensation giving rise to two-cell cleavage stage embryos (Fig 2B).

Oviductal fluid inhibits the heterologous fertilization of bovine ovarian oocytes

We then went on to investigate whether the interaction of oocytes with oviductal fluid would hinder the heterologous fertilization of bovine oocytes with human spermatozoa. In this experiment, oocytes were subjected to 30 minutes of incubation with oviductal fluid derived from the oviducts of slaughtered heifers. Subsequently, these oocytes were incubated with either human or bovine spermatozoa. Results showed that oocytes co-incubated with oviductal fluid could only be fertilized by bovine spermatozoa, with 86% of these oocytes cleaving into two cells. In contrast, bovine oocytes that had been exposed to oviductal fluid were impervious to the penetration by human spermatozoa when using either of the two IVF media (only 3% of these oocytes cleaved into two cells, a similar rate to that recorded in the non-fertilized parthenogenesis group) (see Table S4 and Fig. 2C). These findings suggest that the ZP and/or oolema undergo certain alterations upon contact with oviductal fluid, rendering them penetrable exclusively by spermatozoa of the same species. No sperm were found in the perivitelline space, suggesting that the changes occur in the ZP.

Oviductal fluid sets the species-specificity of the zona pellucida

Earlier version: To validate our hypothesis that the ZP undergoes modifications within the oviduct that confer species-specificity, we used murine oocytes sourced both from the ovary and oviduct. We extracted all the contents inside of the ZP of the oocyte (the cytoplasmic containing all the organelles, nucleus and membranes, and the polar body) (Fig. S4), and then subjected them to either homologous IVF with mouse sperm or heterologous IVF with bovine or human sperm over 4 hours. Because this experiment was done on empty ZPs, we called this test “empty zona penetration test” (EZPT). Our results showed that the ZPs of oocytes obtained directly from the ovary were susceptible to penetration by the sperm of all three species, but once exposed to oviductal fluids, they could only be penetrated by sperm of their homologous species (see Table S5 and Fig. 2D-G). While mouse spermatozoa were found to adhere to the ZP of oocytes obtained both from the ovary and oviduct, more sperm cells appeared on the ZP of oocytes from the ovary. Further, very few bovine and human spermatozoa adhered to the mice ZP of oocytes obtained from the ovary or oviduct (Table S5 and Fig. 2D, E).

Revised version: To validate our hypothesis that the ZP undergoes modifications within the oviduct that confer species-specificity, we used murine oocytes sourced both from the ovary and oviduct. We extracted all the contents inside of the ZP of the oocyte (the cytoplasmic containing all the organelles, nucleus and membranes, and the polar body) (Fig. S4), and then subjected them to either homologous IVF with mouse sperm or heterologous IVF with bovine or human sperm over 4 hours. Because this experiment was done on empty ZPs, we called this test “empty zona penetration test” (EZPT). Our results showed that the ZPs of oocytes obtained directly from the ovary were susceptible to penetration by the sperm of all three species, but once exposed to oviductal fluids, they could only be penetrated by sperm of their homologous species (see Table S5 and Fig. 2D-G). While mouse spermatozoa were found to adhere to the ZP of oocytes obtained both from the ovary and oviduct, more sperm cells appeared on the ZP of oocytes from the ovary. Further, very few bovine and human spermatozoa adhered to the ZP of mouse oocytes obtained from the ovary or oviduct (Table S5 and Fig. 2D, E).

To test whether oviductal fluid indeed determines the species-specificity of bovine oocytes and ascertain whether bovine oviductal fluid affects the ZP, the oocyte, or both, we extracted the oocyte contents, leaving only the ZP (Fig. S4) and exposed the empty bovine ZP to oviductal fluid for 30 minutes. Subsequently, we co-incubated ZP with either bovine, human or murine spermatozoa for a period of 4 hours in an EZPT. As observed in Figure 3, when ZPs were obtained from ovarian oocytes and not exposed to the oviductal fluid, although the numbers of spermatozoa penetrating the ZP were higher in the homologous IVF group, penetration percentages were similar for human, murine or bovine spermatozoa (see Table S6 and Fig. 3A, C, E, G). However, when ZP devoid of oocyte cytoplasmic contents were exposed to oviductal fluid for 30 minutes, these were only susceptible to penetration by bovine spermatozoa (see Table S6 and Fig. 3B, D, F, H), thereby confirming that the effect of oviductal fluid on the ZP is sufficient to determine the specificity of penetrating sperm.

Incubation of empty ZPs obtained from bovine ovarian oocytes with oviductal fluid determines the specificity of spermatozoa capable of penetrating the zona.

The empty zona penetration test (EZPT) in ZPs obtained from IVM bovine oocytes was performed after homologous (bovine sperm) or heterologous (human or murine sperm) fertilization. Similar outcomes were observed when ZPs were not treated with oviductal fluid (A), but after incubation with oviductal fluid for 30 minutes, human and murine sperm were unable to penetrate the bovine ZPs (B) (6 replicates per medium per semen sample). Different letters above error bars (mean ±SD) indicate significant differences (P < 0.05) among groups (ANOVA and Tukey’s post hoc test). Numbers of ZPs used are indicated in Table S6). (C, D, I) Representative pictures of empty bovine ZPs penetrated by bovine sperm illustrating that the penetrating sperm (J) lack an acrosome, while those unable to penetrate the zona maintain the acrosome (K). (E, F, L) Empty bovine ZP penetrated by human sperm revealing that the penetrating sperm have lost the acrosome (M), whereas those not penetrating the zona maintain the acrosome (N). (G, H, O) Representative pictures of empty bovine ZPs penetrated by murine sperm illustrating that the penetrating sperm (P) lack an acrosome, while those unable to penetrate the zona maintain the acrosome (Q). In the absence of oviductal fluid, the empty bovine ZP can be penetrated by bovine (C, I), human (E, L), or murine (G, O) sperm; however, when the ZP has been in contact with oviductal fluid, it can only be penetrated by bovine sperm (D), and not by human (F) or mouse (H) sperm. Scale bar for ZP pictures = 50 µm. Scale bar for sperm pictures = 5 µm.

With our EZPT, we were able to determine that all spermatozoa penetrating the ZP lacked an acrosome, while the majority of those that failed to penetrate the egg coat retained an intact acrosome (see Fig. 3I-Q). To assess acrosomal status, the fluorescein isothiocyanate-conjugated peanut agglutinin (FITC-PNA) method was used (38). This observation suggests that during the penetration of the ZP, bovine, murine, and human spermatozoa undergo the acrosomal reaction. Additionally, it should be noted that, regardless of incubation with oviductal fluid, a higher proportion of bovine spermatozoa remained bound to the ZP (see Table S6 and Fig. 3C, D, I). It may be seen in Table S6 and Figure 3A-B, that more bovine spermatozoa bound to the ZP than human and murine ones. No difference was found between numbers of bovine sperm bound to ZP treated or not with oviductal fluid. Also, for human spermatozoa, no difference was found between those bound to ZPs treated or not with oviductal fluid, but for mice sperm, a greater number attached to the bovine ZP when it had not been in contact with oviductal fluid.

Earlier version: Homology of bovine, human and murine OVGP1 proteins and fusion of recombinant OVGP1 proteins with murine or bovine ZPs

Revised version: Homology of bovine, human and murine OVGP1 proteins and binding of recombinant OVGP1 proteins with murine or bovine ZPs

Earlier version: It has been reported that OVGP1 in pigs influences sperm binding to the ZP in a species-specific manner (39) and that this protein facilitates specific sperm binding to the ZP, potentially contributing to species-specificity (29, 33). To investigate the role of OVGP1 on the species-specificity of the ZP, we have use histidine-tagged cow recombinant OVGP1 glycoproteins and DDK-tagged human and mouse recombinant OVGP1 glycoprotein. As depicted in Fig. 4A, the sequences of these three OVGP1 have five distinct regions (A, B, C, D and E) (29). Regions A and C are conserved in the different mammals. Region A, corresponding to the amino terminal end, shows high identity among monotremes, marsupials and placentals. Region B features low identity among the different mammals and contains multiple insertions/deletions. Region D is typical of human oviductin and region E is an insertion present only in the mouse (Mus) . Details of OVGP1 sequence alignments for the three species can be seen in Figure S5. Clear differences may be seen in predicted N-glycosylation sites and chitin-binding sites (Fig. S5) as well as in predicted O-glycosylation sites (Table S7).

Revised version: It has been reported that OVGP1 in pigs influences sperm binding to the ZP in a species-specific manner (39) and that this protein facilitates specific sperm binding to the ZP, potentially contributing to species-specificity (29, 33). To investigate the role of OVGP1 on the species-specificity of the ZP, we have use histidine-tagged cow recombinant OVGP1 glycoproteins and DDK-tagged human and mouse recombinant OVGP1 glycoprotein. As depicted in Fig. 4A, the sequences of these three OVGP1 have five distinct regions (A, B, C, D and E) (29). Regions A and C are conserved in the different mammals. Region A, corresponding to the amino terminal end, shows high identity among monotremes, marsupials and placentals. Region B features low identity among the different mammals and contains multiple insertions/deletions. Region D is typical of human oviductin and region E is an insertion present only in the mouse (Mus). Details of OVGP1 sequence alignments for the three species can be seen in Figure S5 (35, 40, 41) The role of oviduct-specific glycoprotein (OVGP1) in Clear differences may be seen in predicted N-glycosylation sites and chitin-binding sites (Fig. S5) as well as in predicted O-glycosylation sites (Table S7).

Structure of oviductin, Western blots of OVGP1 recombinants and localization of these recombinants at the bovine or murine ZPs.

Earlier version: (A) Diagram showing the five regions (A, B, C, D and E) present in some of the oviductin proteins of human (hOVGP1), murine (mOVGP1), and bovine (bOVGP1) mammalian species (Figure taken from (29)). (B) Western blots of the three OVGP1 recombinants proteins used in this study (human, murine, and bovine). Proteins were expressed in mammalian cells, separated by SDS-PAGE and analyzed by immunoblotting using rabbit polyclonal antibody to the human OVGP1. Oviductal fluid of mice in oestrus and anoestrus indicating the presence in estrous of the OVGP1 band. Oviductal fluid from ovulated cows or anoestrus cows. ZPs from IVM murine (C) and bovine (D) oocytes were incubated for 30 minutes at RT with recombinants bOVGP1, mOVGP1, and hOVGP1. ZPs were fixed and imaged by confocal fluorescence and DIC microscopy using rabbit polyclonal antibody to the human OVGP1 for bOVGP1, and a monoclonal antibody against Flag-tag for hOVGP1 and mOVGP1. Scale bars = 20 μm.

Revised version: (A) Diagram showing the five regions (A, B, C, D and E) present in some of the oviductin proteins of human (hOVGP1), murine (mOVGP1), and bovine (bOVGP1) mammalian species (Figure adapted from (29)). (B) Western blots of the three OVGP1 recombinants proteins used in this study (human, murine, and bovine). Recombinant bovine protein was expressed in BHK-21 cells and purified; whereas recombinant murine and human proteins were purchased by Origene and had been produced in HEK293T. Then, proteins were separated by SDS-PAGE and analyzed by immunoblotting using rabbit polyclonal antibody to the human OVGP1. The following lanes of the gel contain the protein extracts from oviductal fluid of female mice in oestrus and anoestrus and from oviductal fluid of ovulated cows or anoestrus cows, indicating with an arrow the presence in estrous of the OVGP1 band for both species. ZPs from IVM murine (C) and bovine (D) oocytes were incubated for 30 minutes at RT with recombinants bOVGP1, mOVGP1, and hOVGP1. ZPs were fixed and imaged by confocal fluorescence and DIC microscopy using rabbit polyclonal antibody to the human OVGP1 for bOVGP1, and a monoclonal antibody against Flag-tag for hOVGP1 and mOVGP1. Scale bars = 20 μm.

Earlier version: Proteins were expressed in mammalian cells, separated by SDS-PAGE and analyzed by immunoblotting using rabbit polyclonal antibody to human OVGP1. Additionally, rabbit monoclonal antibody against His-tag, and mouse monoclonal antibody against Flag (DDK)-tag immunoblots were performed to test the recombinant OVGP1 proteins (data not shown). Western blots of the three OVGP1 recombinants indicated expected sizes based on those of the proteins: 75 kDa for human and murine OVGP1 and around 60 kDa for bovine OVGP1 (Fig. 4B and S6). Immunoblots were prepared of bovine and murine oviductal fluid protein extracts from animals in oestrus (mice) or ovulated (cow) versus anestrous animals. OVGP1 in oviductal fluids was in the range of 100-150 ng/µl in mice and 150-200 ng/µL in cow. In Figure 4B, an arrow indicates the presence of the expected band of OVGP1 found in the oviductal fluid of female mice in oestrus or oviduct of ovulated cows, compared to the oviductal fluid of animals in anestrous. Murine oviductal fluids displayed the expected 75kDa band and a more intense 50kDa band, but only the 75kDa band corresponded to OVGP1 glycoprotein, as confirmed by proteomics of the bands along with PEAKS Studio v11.5 search engine peptide identification software.

Revised version: Proteins were expressed in mammalian cells, separated by SDS-PAGE and analyzed by immunoblotting using rabbit polyclonal antibody to human OVGP1. Additionally, rabbit monoclonal antibody against His-tag, and mouse monoclonal antibody against Flag (DDK)-tag immunoblots were performed to test the recombinant OVGP1 proteins (data not shown). Western blots of the three OVGP1 recombinants indicated expected sizes based on those of the proteins: 75 kDa for human and murine OVGP1 and around 60 kDa for bovine OVGP1 (Fig. 4B and S6). Immunoblots were prepared of bovine and murine oviductal fluid protein extracts from animals in oestrus (mice) or ovulated (cow) versus anoestrus animals. OVGP1 in oviductal fluids was in the range of 100-150 ng/µl in mice and 150-200 ng/µL in cow. In Figure 4B, an arrow indicates the presence of the expected band of OVGP1 found in the oviductal fluid of female mice in oestrus or oviduct of ovulated cows, compared to the oviductal fluid of animals in anoestrus. Murine oviductal fluids displayed the expected 75kDa band and a more intense 50kDa band, but only the 75kDa band corresponded to OVGP1 glycoprotein, as confirmed by proteomics of the bands along with PEAKS Studio v11.5 search engine peptide identification software. As according to Proteomic results, in the 75kDa band OVGP1 was identified with 14 unique peptides over 240 protein groups, whereas in the 50kDa band it was not present OVGP1 over 307 identified protein groups.

OVGP1 recombinant proteins were also assessed for their ability to bind murine or bovine ZPs obtained from in vitro matured (IVM) ovarian oocytes compared to ZPs derived from murine oocytes from superovulated females (Fig. 4 and Fig S7). Oocytes were incubated for 30 minutes at room temperature with bovine (bOVGP1), murine (mOVGP1), or human (hOVGP1) OVGP1, and were then fixed and imaged by confocal fluorescence and DIC microscopy using rabbit polyclonal antibody to human OVGP1 for bOVGP1 and endogenous OVGP1, and mouse monoclonal antibody against Flag (DDK)-tag for hOVGP1 and mOVGP1. All three recombinant proteins were able to bind to the murine and bovine ZPs (Fig. 4C-D; controls in Fig. S7). No differences were observed in the ability of the three OVGP1s to attach and penetrate the mouse ZPs (Fig. 4C and Fig S7A). However, while binding of bOVGP1 was observed throughout the thickness of the bovine ZP, mOVGP1 and hOVGP1 preferentially bound to the outer bovine ZP surface without penetrating the coat (Fig. 4D and Fig S7B).

OVGP1 plays a species-specific role in setting the specificity of the zona pellucida

To investigate whether oviductal protein OVGP1 had a similar effect to oviductal fluid in determining the species-specificity of bovine ZP during fertilization, we conducted the EZPT using ZPs from ovarian bovine oocytes that were either treated with bOVGP1 or left untreated (for 30 minutes). In this experiment, we used homologous (bovine) and heterologous sperm (human and murine). As illustrated in Figure 5, when the ZPs from ovarian oocytes had not been exposed to bovine OVGP1, both homologous and heterologous sperm could penetrate the ZPs (Table S8 and Fig. 5A). However, after ZPs were treated with bOVGP1 for 30 minutes, only bovine sperm were capable of penetrating the ZP, while human or mouse sperm were unable to do this (Table S8 and Fig. 5B). Similar effects were found in terms of numbers of penetrated sperm per ZP (Fig 5A, B).

Incubation of empty ZPs obtained from bovine ovarian oocytes with OVGP1 determines the specificity of spermatozoa capable of penetrating the zona.

The empty zona penetration test (EZPT) was performed after homologous (bovine sperm) or heterologous (human or murine sperm) fertilization. Similar penetration rates were observed when ZPs were not treated with bOVGP1 protein (A), but after incubation with bOVGP1 for 30 minutes, human or murine sperm were unable to penetrate the bovine ZPs (B). A drastic reduction was also observed for fertilization by heterologous sperm, both in the average number of sperm penetrating the ZPs, and in the average of number of sperm binding to the ZPs, when fertilization without pretreatment of the ZP with bOVGP1 (A) was compared to fertilization after the zona had been in contact with bOVGP1 (B). (6 replicates per medium per semen sample). Different letters above error bars (mean ±SD) indicate significant differences (P < 0.05) among groups (ANOVA and Tukey’s post hoc test). Numbers of ZPs used are indicated in Table S8).

Earlier version: ZPs that had not been exposed to bovine OVGP1 all showed attached sperm of the three species; bovine sperm sticking most, followed by mouse sperm and then human sperm (Fig. 5A). However, when ZPs had been exposed to bovine OVGP1, only bovine sperm were found to bind to bovine ZPs (Fig. 5B).

Revised version: ZPs that had not been exposed to bovine OVGP1 all showed attached sperm of the three species; bovine sperm specific binding most, followed by mouse sperm and then human sperm (Fig. 5A). However, when ZPs had been exposed to bovine OVGP1, only bovine sperm were found to bind to bovine ZPs (Fig. 5B).

Earlier version: Due to that under in vivo conditions sperm comes into contact with the oviductal tract of the female, and consequently with OVGP1, we studied the potential effect of this oviductal fluid protein on the behaviour of the semen sample when exposed to ZPs. To this end, we conducted the EZPT using ZPs from ovarian bovine oocytes that were either treated with bOVGP1 or left untreated (for 30 minutes). In this experiment, we used both homologous (bovine) and heterologous (murine) sperm. Additionally, capacitated sperm was either treated or left untreated for 30 minutes with bovine or murine OVGP1. As shown in Table S9, we found that treating semen samples with bovine or murine OVGP1 did not affect their ability to penetrate or bind to ZPs, yielding results similar to those observed in previous experiments (Fig 5A,B and Table S8).

Revised version: Since, under in vivo conditions, sperm comes into contact with the female’s oviductal tract, and consequently with OVGP1, we investigated the potential effect of this oviductal fluid protein on the behavior of the semen sample when exposed to ZPs. To this end, we conducted the EZPT using ZPs from ovarian bovine oocytes that were either treated with bOVGP1 or left untreated (for 30 minutes). In this experiment, we used both homologous (bovine) and heterologous (murine) sperm. Additionally, capacitated sperm was either treated or left untreated for 30 minutes with bovine or murine OVGP1. As shown in Table S9, we found that treating semen samples with bovine or murine OVGP1 did not affect their ability to penetrate or bind to ZPs, yielding results similar to those observed in previous experiments (Fig 5A,B and Table S8).

Earlier version: Scanning electron microscopy (SEM) observation of the effect of oviductal fluid and homologous or heterologous OVGP1 on bovine zona pellucida structure

Revised version: Observation by scanning electron microscopy (SEM) of the effect of oviductal fluid and homologous or heterologous OVGP1 on the structure of the bovine zona pellucida

After in vitro maturation (IVM), oocytes displayed a porous ZP structure characterized by fine-meshed reticular pores; the pores were circular or elliptical, arbitrarily distributed, and accompanied by deep pore holes (Fig. 6A-B). Upon exposure to bovine OVGP1 for 30 minutes, bOVGP1 deposition on the outer ZP surface was observed (Fig. 6C-D). Pores were larger and more numerous for IVM oocytes that had not been in contact with bOVGP1. A similar effect was observed with bOF (Fig. S8A-B), and with human or murine OVGP1 (Fig. S8C-D). In this SEM study, we were thus able to visualize the morphological changes occurring on the outer ZP surface of IVM bovine oocytes following their exposure to OVGP1.

Scanning electron micrographs (SEM) of the outer surface of the bovine ZP treated or not with OVGP1.

The ZP of an in vitro matured (IVM) bovine oocyte: Magnification x2000 (A) and x4000 (B); and bovine oocytes IVM in the presence of bovine OVGP1 (bOVGP1): Magnification x2000 (C) and x4000 (D). High magnification reveals the ultrastructural characteristics of the ZP’s pores on the IVM oocytes without bOVGP1 (B) and with bOVGP1 (D). Scale bars = 10 μm.

Earlier version: Only species-specific OVGP1 confers the zona pellucida complete species-specificity

Revised version: Only species-specific OVGP1 confers species-specificity to the zona pellucida

Earlier version: We conducted an experiment to determine whether oviductin from a species different from the origin of the ZP would confer specificity. Through the EZPT several combinations were tested whereby ZPs from bovine or murine ovarian oocytes were co-incubated with bovine, murine, or human OVGP1 and then exposed to sperm of the three species. As shown in Figure 7A and Table S10, for bovine ZPs and bOVGP1, penetration of the empty ZP was only possible in the case of bovine sperm, such that the entry of sperm of the other two species was completely blocked. For the combination bovine ZP/mOVGP1, human sperm were fully blocked, while some penetration (approximately 60%) by bovine sperm (homologous to the ZP) or murine sperm (homologous to OVGP1) was possible. Similar results were observed for bovine ZP/hOVGP1, sperm homologous to the ZP (bovine) and their pretreatment with human OVGP1, giving rise to partial penetration capacity (60%), indicating that OVGP1 non-homologous to the ZP allow the partial passage of sperm of the ZP’s species of origin and of the OVGP1-providing species. Similar penetration results were obtained using murine ZPs in combination with the three oviductin and sperm species (Fig. 7B, Table S11), indicating that results were not specific to bovine ZP.

Revised version: We conducted an experiment to determine whether oviductin from a species different from the origin of the ZP could confer specificity. Through the EZPT several combinations were tested whereby ZPs from bovine or murine ovarian oocytes were co-incubated with bovine, murine, or human OVGP1 and then exposed to sperm of the three species. As shown in Figure 7A and Table S10, for bovine ZPs and bOVGP1, penetration of the empty ZP was only possible in the case of bovine sperm, such that the entry of sperm of the other two species was completely blocked. For the combination bovine ZP/mOVGP1, human sperm were fully blocked, while some penetration (approximately 60%) by bovine sperm (homologous to the ZP) or murine sperm (homologous to OVGP1) was possible. Similar results were observed for bovine ZP/hOVGP1, sperm homologous to the ZP (bovine) and their pretreatment with human OVGP1, giving rise to partial penetration capacity (60%), indicating that OVGP1 non-homologous to the ZP allow the partial passage of sperm of the ZP’s species of origin and of the OVGP1-providing species. Similar penetration results were obtained using murine ZPs in combination with the three oviductin and sperm species (Fig. 7B, Table S11), indicating that results were not specific to bovine ZP.

Species-specific OVGP1 confers complete species-specificity to the zona pellucida.

Experiment using the EZPT to analyze several combinations whereby ZPs from bovine and murine ovarian IVM oocytes were co-incubated with bovine, murine, or human oviductin, and exposed to sperm of all three species. The variables ZP penetration rate, average number of sperm penetrating the ZPs, and average number of sperm bound to ZPs were examined using bovine ZPs (A) or murine ZPs (B) and combinations of one of the three oviductins with sperm of one of the three species. Only by combining ZP, OVGP1 and sperm of the same species, was species-specific fertilization possible. However, when the same species was matched only with the oviductin and sperm, or only with the ZP and sperm, penetration and binding to the ZP of the spermatozoa also occurred, although in much smaller measure. No penetration was observed when only the ZP and OVGP1 species matched. Different letters above error bars (mean ±SD) indicate significant differences (P < 0.05) among groups (ANOVA and Tukey’s post hoc test). Numbers of ZPs used are indicated in Table S10-11).

Results corresponding to mean numbers of penetrated sperm (Fig. 7A and B, center graphs) were consistent with these previous findings reflecting a clear reduction when the OVGP1 source was not homologous to that of the ZP, but the sperm used belonged to one of these two species. When there was no species homology among ZP, OVGP1, or sperm, the mean number of sperm recorded was less than one.

When we examined mean numbers of sperm bound to the ZP (Fig. 7A and B, right graphs), results paralleled those of penetration. Hence, when there was homology among the sperm, ZP, and OVGP1, only sperm of the homologous species were able to bind to the ZP; and when there was only homology between sperm and either OVGP1 or ZP, there was a reduction in the number of sperm found attached to the ZP. Interestingly, species homology between OVGP1 and sperm resulted in more bound sperm than when there was only homology between ZP and sperm, indicating the involvement of OVGP1 in sperm-ZP binding. These differences in sperm binding were, nevertheless, not reflected in the penetration rate.

Impact of neuraminidase on the EZPT in bovine and murine oocytes pre and post OVGP1 treatment

Earlier version: The enzyme neuraminidase (NMase) cleaves sialic acid residues commonly found in glycoproteins and glycolipids, and is thus able to modify the surface properties of cells and consequently affect their interactions with other molecules or cells. While the involvement of sialic acid in sperm–ZP binding (28, 40) is generally accepted, studies reaching this conclusion have not compared the ZPs of oocytes that have been in contact or not with OVGP1. The aim of this experiment was to examine the influence of NMase treatment on the sperm’s ability to bind and to penetrate the ZPs of both bovine and murine oocytes before and after exposure to OVGP1 (Fig. 8, Tables S12 and S13).

Revised version: The enzyme neuraminidase (NMase) cleaves sialic acid residues commonly found in glycoproteins and glycolipids, and is thus able to modify the surface properties of cells and consequently affect their interactions with other molecules or cells. While the involvement of sialic acid in sperm–ZP binding (28, 42) is generally accepted, studies reaching this conclusion have not compared the ZPs of oocytes that have been in contact or not with OVGP1. The aim of this experiment was to examine the influence of NMase treatment on the sperm’s ability to bind and to penetrate the ZPs of both bovine and murine oocytes before and after exposure to OVGP1 (Fig. 8, Tables S12 and S13).

Effect of neuraminidase (NMase) treatment of bovine or murine ZPs before and after contact with OVGP1 on sperm penetration.

Penetration and sperm binding rates were measured after homologous EZPT with bovine sperm using ZPs not subjected (A) or subjected (B) to 30 minutes of incubation with bOVGP1. Penetration and sperm binding rates were measured after homologous EZPT with murine sperm using ZPs not subjected (C) or subjected (D) to 30 minutes of incubation with mOVGP1. ZPs of both species were co-incubated with acetate buffer (PH 4.5) at 38°C for 18 hours in the presence or absence of neuraminidase diluted at 5 UI/mL. Different letters above error bars (mean ±SD) indicate significant differences (P < 0.05) among groups (ANOVA and Tukey’s post hoc test). Numbers of ZPs used are indicated in Table S12-13).

The results indicate that, regardless of whether ZPs had been or not in contact with OVGP1, after NMase treatment, sperm penetration was completely blocked, and the number of sperm adhering to the ZP was reduced by 70%, suggesting that both nonspecific sperm-ZP binding and species-specific binding occurring after contact with OVGP1 are prevented by NMase treatment, thus impeding sperm penetration (Fig. 8, Tables S12 and S13). It is important to note that despite completely blocking penetration, the binding of some sperm to the ZP likely indicates their ability to nonspecifically adhere, which could occur, for example, when the acrosome is damaged.

To confirm that the sialic acids on bOVGP1 are responsible for the species-specific binding of sperm, we conducted an experiment in which only OVGP1 was treated with NMase before coming into contact with the bovine ZP. The treated bOVGP1 was then incubated with the ZP for 30 min, and then we performed the homologous (bovine sperm) or heterologous (mouse sperm) EZPT. In this setup, the ZP retained its sialic acid residues, while OVGP1 did not. As shown in Figure 9 and Table S14, the treatment did not reduce the percentage of penetration when homologous sperm was used, but it halved the number of sperm penetrating the ZP (Fig. 9A). It is possible that penetration still occurs because there is sialic acid on the ZP, which allows non-species-specific sperm binding. However, the number of sperm that penetrate is lower because there is no sialic acid on OVGP1, blocking species-specific binding.

Effect of neuraminidase (NMase) treatment of OVGP1 on EZPT using bovine ZPs with homologous or heterologous sperm.

(A) Penetration rate, average number of sperm within penetrated ZPs, and average of number of sperm bound to ZPs were measured after homologous EZPT with bovine sperm, using ZPs untreated with NMase and bOVGP1 either treated or untreated with NMase. (B) Penetration rate, average number of sperm within penetrated ZPs, and average of number of sperm bound to ZPs were measured after homologous (bovine sperm) or heterologous (mouse sperm) EZPT, using ZPs untreated with NMase, either without bOVGP1 or with bOVGP1 treated or untreated with NMase. Numbers of ZPs used are indicated in Table S14

Unlike what happens with bovine sperm, penetration occurred with heterologous mouse sperm in the absence of bOVGP1. However, no penetration was observed in intact bovine ZPs in the presence of non treated bOVGP1, likely because bOVGP1 and its bovine-specific sialic acids prevent the binding and penetration of mouse sperm. When OVGP1 was treated with NMase, a certain degree of penetration was observed, likely due to the sialic acids present on the ZP, which enable non-species-specific sperm binding (Fig. 9B). These findings highlight the critical role of sialic acids in sperm binding, with those on the ZP facilitating non-specific binding and those on OVGP1 playing a key role in species-specific binding.

Added in the revised version: Discussion

Added in the revised version: The zona pellucida (ZP) exhibits different characteristics before and after ovulation and fertilization (3). The present study provides evidence that both bovine and murine oocytes acquire species-specific fertilization capacity within the oviduct. Oocytes from the ovaries of cows and mice can be fertilized by sperm from distant mammalian species, yet they acquire species-specificity upon contact with oviductal fluid or specific OVGP1 protein (Fig. 10). Our results unveil two distinct mechanisms for sperm to recognize and penetrate the ZP. One is nonspecific and acts when oocytes come directly from the ovary, while the other is specific and operates when oocytes reach the oviduct and encounter oviductal fluid or its main secreted constituent, OVGP1. In the nonspecific mechanism, carbohydrate residues present in the ZP, particularly those containing sialic acid, are necessary, allowing sperm to penetrate through the ZP’s matrix. For the species-restricted mechanism, oligosaccharides chains containing sialic acids from OVGP1 must be recognized by sperm for them to bind and penetrate the ZP. These findings indicate that species-specific OVGP1 determines the species-specificity of the ZP in mammals. Furthermore, our results suggest that heterologous IVF using bovine or murine ovarian oocytes or their ZPs could be an excellent approach for assessing sperm capacitation and fertility in mammalian species where obtaining same-species oocytes is challenging. This study contributes to our understanding of reproductive biology and sheds light on the molecular mechanisms driving cross-species fertilization.

Model. Schematic representation of the heterologous fertilization model via EZPT under multiple conditions.

Earlier version: The figure illustrates how sperm penetration into the zona pellucida (ZP) occurs non-specifically among various mammalian species when the ZP has not been exposed to homologous oviductal fluid or OVGP1 (left diagram). However, this interaction becomes species-specific when the ZP is incubated with oviductal fluid or OVGP1 from the same species (central diagram); being this specificity lost when OVGP1 is heterologous to the ZP (right diagram).

Revised version: The figure illustrates how sperm penetration into the zona pellucida (ZP) occurs non-specifically among various mammalian species when the ZP has not been exposed to homologous oviductal fluid or OVGP1 (left diagram). However, this interaction becomes species-specific when the ZP is incubated with oviductal fluid or OVGP1 from the same species (central diagram) This specificity is lost when OVGP1 is heterologous to the ZP (right diagram).

Removed in the revised version: Discussion

Removed in the revised version: The zona pellucida (ZP) exhibits different characteristics before and after ovulation and fertilization (3). The present study provides evidence that both bovine and murine oocytes acquire species-specific fertilization capacity within the oviduct. Oocytes from the ovaries of cows and mice can be fertilized by sperm from distant mammalian species, yet they acquire species-specificity upon contact with oviductal fluid or specific OVGP1 protein. Our results unveil two distinct mechanisms for sperm to recognize and penetrate the ZP. One is nonspecific and acts when oocytes come directly from the ovary, while the other is specific and operates when oocytes reach the oviduct and encounter oviductal fluid or its main secreted constituent, OVGP1. In the nonspecific mechanism, carbohydrate residues present in the ZP, particularly those containing sialic acid, are necessary, allowing sperm to penetrate through the ZP’s matrix. For the species-restricted mechanism, oligosaccharides chains containing sialic acids from OVGP1 must be recognized by sperm for them to bind and penetrate the ZP. These findings indicate that species-specific OVGP1 determines the species-specificity of the ZP in mammals. Furthermore, our results suggest that heterologous IVF using bovine or murine ovarian oocytes or their ZPs could be an excellent approach for assessing sperm capacitation and fertility in mammalian species where obtaining same-species oocytes is challenging. This study contributes to our understanding of reproductive biology and sheds light on the molecular mechanisms driving cross-species fertilization.

Species-specific interactions between gamete membranes in mammals are not always clear-cut. While the exact mechanism of fusion between sperm and oolemma membranes is only partially understood, several molecules such as IZUMO, TMEM95, SPACA6, and SOF1 have been incriminated (6–8). The empty zona penetration test (EZPT) enables heterologous sperm to overcome the oocyte’s second barrier, the plasma membrane or oolema. The test known as HEPT, which employs ZP-free oocytes, has shown that the hamster oocyte membrane can fuse with the sperm membrane of many mammalian species but the presence of ZP allows only the penetration of hamstersperm. Our study reveals that bovine and murine ovarian oocytes permit the penetration of both the ZP and oolema of various mammalian orders. This suggests that, prior to reaching the oviduct, a similar mechanism is used across mammalian species for heterologous sperm to recognize and penetrate the ZP and oocyte plasma membrane. However, when the oviduct is reached, some components play a crucial role in inducing species-specific alterations in the ZP, preventing the recognition and/or penetration of the sperm of other species.

Earlier version: During folliculogenesis and transit through the oviduct, the ZP of ovarian oocytes undergoes maturation in preparation for interactions with spermatozoa in the female reproductive tract (36, 41, 42). These changes occur when ovulated oocytes are exposed to oviductal fluid, although the exact molecular mechanisms are incompletely understood, it was previously reported changes in carbohydrate composition and incorporation of proteins (41, 43). It has been observed that immature oocytes possess a smooth, compact ZP surface with few pores, while IVM oocytes exhibit a net-like porous surface with a rough pattern on the zona matrix (44, 45). Here, we observed that IVM ZPs exposed to OVGP1, undergo a complete transformation, oviductin covering most of the ZP and its pores. This suggests that

Revised version: During folliculogenesis and transit through the oviduct, the ZP of ovarian oocytes undergoes maturation in preparation for interactions with spermatozoa in the female reproductive tract (36, 43, 44). These changes occur when ovulated oocytes are exposed to oviductal fluid, although the exact molecular mechanisms are incompletely understood, it was previously reported changes in carbohydrate composition and incorporation of proteins (43, 45). It has been observed that immature oocytes possess a smooth, compact ZP surface with few pores, while IVM oocytes exhibit a net-like porous surface with a rough pattern on the zona matrix (46, 47). Here, we observed that IVM ZPs exposed to OVGP1, undergo a complete transformation, oviductin covering most of the ZP and its pores. This suggests that

Earlier version: OVGP1 contributes to the maturation of the oviductal ZP by influencing matrix morphology. At the point of fertilization, OVGP1 is found as the main non-serum glycoprotein in oviductal fluid and its relationship with the ZP was noted early on (46). OVGP1 plays a role in strengthening the ZP, rendering it resistant to proteolytic digestion and sperm penetration (36), and these processes are regulated by OVGP1’s C-terminal region (47). The C-terminal region of OVGP1 is the region which mainly differ among species (29) and this fact may explain the species-specificity of sperm penetrability. A recently study demonstrates that the OVGP1 gene is essential for normal early embryonic development, as OVGP1-KO hamsters exhibit fertilization failures, abnormal embryos, and low implantation rates. These findings highlight the critical role of the oviductal microenvironment in regulating early embryogenesis, however, they did not analyze the interaction between OVGP1 and ZP (48).

Revised version: OVGP1 contributes to the maturation of the oviductal ZP by influencing matrix morphology. At the point of fertilization, OVGP1 is found as the main non-serum glycoprotein in oviductal fluid and its relationship with the ZP was noted early on (48). OVGP1 plays a role in strengthening the ZP, rendering it resistant to proteolytic digestion and sperm penetration (36), and these processes are regulated by OVGP1’s C-terminal region (49). The C-terminal region of OVGP1 is the region which mainly differ among species (29) and this fact may explain the species-specificity of sperm penetrability. A recently study demonstrates that the OVGP1 gene is essential for normal early embryonic development, as OVGP1-KO hamsters exhibit fertilization failures, abnormal embryos, and low implantation rates. These findings highlight the critical role of the oviductal microenvironment in regulating early embryogenesis, however, they did not analyze the interaction between OVGP1 and ZP (50).

Earlier version: In mammals, it is accepted that fertilization typically occurs within the ampulla of the oviduct, where sperm bind to the oocyte through interactions with ZP via N- and O-glycoproteins (49). Over time, the proposed molecular model explaining sperm-oocyte interactions has evolved. Initially, studies suggested ZP2 and ZP3 acted as primary sperm receptors (50–52). Later, it was proposed that sperm-ZP interactions might depend more on the supramolecular structure of the oocyte coat rather than specific ZP subunits (53). More recent research points to a nuanced perspective, highlighting a central role for a specific domain of ZP2 in gamete recognition, which physically interacts with complementary molecules on sperm (54). Despite these advances, the exact roles of ZP and OVGP1 glycosylation and sialic acid residues (Sias) in fertilization are poorly understood. After ovulation, the changes reported in the carbohydrate composition of the ZP (3, 28) are likely induced by the addition of glycoproteins of oviductal origin, as we have seen here with OVGP1. By employing NMase, to remove sialic acid residues from ZP glycoproteins, we here ascertain that OVGP1 plays a pivotal role in the modifications to the carbohydrate composition of the ZP necessary for species-specific sperm fertilization. In conclusion, our results indicate that the oocyte ZP must interact with the contents of the oviductal fluid, specifically with OVGP1 protein, to ensure that only sperm from the same species can fertilize the egg. This determines that the modifications introduced by OVGP1 in the ZP are critical to define the barrier among the different species of mammals. Additionally, future studies would be valuable to investigate whether incubating oocytes with oviductal fluid (or OVGP1) could reduce polyspermy in porcine IVF and whether ZPs could be leveraged to naturally enhance sperm selection in human ICSI.

Revised version: In mammals, it is accepted that fertilization typically occurs within the ampulla of the oviduct, where sperm bind to the oocyte through interactions with ZP via N- and O-glycoproteins (51). Over time, the proposed molecular model explaining sperm-oocyte interactions has evolved. Initially, studies suggested ZP2 and ZP3 acted as primary sperm receptors (52–54). Later, it was proposed that sperm-ZP interactions might depend more on the supramolecular structure of the oocyte coat rather than specific ZP subunits (55). More recent research points to a nuanced perspective, highlighting a central role for a specific domain of ZP2 in gamete recognition, which physically interacts with complementary molecules on sperm (56). Despite these advances, the exact roles of ZP and OVGP1 glycosylation and sialic acid residues (Sias) in fertilization are poorly understood. After ovulation, the changes reported in the carbohydrate composition of the ZP (3, 28) are likely induced by the addition of glycoproteins of oviductal origin, as we have seen here with OVGP1. By employing NMase, to remove sialic acid residues from ZP glycoproteins, we here ascertain that OVGP1 plays a pivotal role in the modifications to the carbohydrate composition of the ZP necessary for species-specific sperm fertilization. In conclusion, our results indicate that the oocyte ZP must interact with the contents of the oviductal fluid, specifically with OVGP1 protein, to ensure that only sperm from the same species can fertilize the egg. This determines that the modifications introduced by OVGP1 in the ZP are critical to define the barrier among the different species of mammals. It is noteworthy that this work strengthens the idea that incubating oocytes with oviductal fluid (or OVGP1) can be used to reduce polyspermy in porcine IVF (36, 57, 58). Additionally, future studies would be valuable to investigate if ZPs could naturally enhance sperm selection in human ICSI.

Materials and methods

Reagents

Earlier version: All media components were purchased from Sigma–Aldrich, except where stated otherwise.

Revised version: All media components were purchased from Sigma–Aldrich(St. Louis, MO, USA), except where stated otherwise.

Sperm collection, cryopreservation, thawing, and capacitation

Earlier version: Human semen samples were obtained from ejaculated sperm samples of three normozoospermic donors with their written informed consent. Approval for the study protocol was obtained from the Hospital Clinico San Carlos Research Ethics Review Committee (Madrid, Spain) and the FivCenter fertility clinic (Aravaca, Madrid, Spain) in accordance with the principles of the Declaration of Helsinki. Informed consent for patients was overseen by the Spanish Fertility Society. Donor semen samples were obtained by masturbation after 2-3 days of abstinence and analyzed in the fertility clinic’s laboratory. Samples were left to liquefy and then subjected to basic seminogram analysis following WHO guidelines (55). Sperm concentration and motility assessments were conducted using a Makler® counting chamber (BioCare Europe, Rome, Italy).

Revised version: Human semen samples were obtained from ejaculated sperm samples of three normozoospermic donors with their written informed consent. Approval for the study protocol was obtained from the Hospital Clinico San Carlos Research Ethics Review Committee (Madrid, Spain) and the FivCenter fertility clinic (Aravaca, Madrid, Spain) in accordance with the principles of the Declaration of Helsinki. Informed consent for patients was overseen by the Spanish Fertility Society. Donor semen samples were obtained by masturbation after 2-3 days of abstinence and analyzed in the fertility clinic’s laboratory. Samples were left to liquefy and then subjected to basic seminogram analysis following WHO guidelines (59). Sperm concentration and motility assessments were conducted using a Makler® counting chamber (BioCare Europe, Rome, Italy).

Earlier version: Semen samples were diluted at a 1:1 ratio (vol:vol) with the freezing medium (TEST Yolk Buffer Fujifilm, Irvine Scientific) and stored for 40 min at 4°C for equilibration. After this time, the mixture was plunged drop by drop (10µl) onto a dry ice plate, so that small solid microspheres were formed. These microspheres were then recovered and placed into cryovials for storage in liquid nitrogen at -196°C. For thawing, the desired pellets were transferred to a 15 mL Falcon tube and kept at 37°C for 20 minutes. Next, the samples were washed by diluting 1:1 in human G-IVF™ PLUS medium (Vitrolife, Igenomix) and centrifuging at 1800 rpm for 10 minutes. After discarding the supernatant, 200-220 µL of G-IVF™ PLUS medium was added to the sediment and the sample stored at 38.5°C in a 5% CO2 atmosphere for 45-60 minutes. After this time, 90 µL of the upper culture medium layer was removed using a sterile pipette. This aliquot was ready for use following standard swim-up procedures.

Revised version: Semen samples were diluted at a 1:1 ratio (vol:vol) with the freezing medium (TEST Yolk Buffer Fujifilm, Irvine Scientific) and stored for 40 min at 4°C for equilibration. After this time, the mixture was plunged drop by drop (10µl) onto a dry ice plate, so that small solid microspheres were formed. These microspheres were then recovered and placed into cryovials for storage in liquid nitrogen at -196°C. For thawing, the desired pellets were transferred to a 15 mL Falcon tube and kept at 37°C for 20 minutes. Next, the samples were washed by diluting 1:1 in human G-IVF™ PLUS medium (Vitrolife, Igenomix) and centrifuging at 350g for 10 minutes. After discarding the supernatant, 200-220 µL of G-IVF™ PLUS medium was added to the sediment and the sample stored at 38.5°C in a 5% CO2 atmosphere for 45-60 minutes. After this time, 90 µL of the upper culture medium layer was removed using a sterile pipette. This aliquot was ready for use following standard swim-up procedures.

Earlier version: Sperm from B6CBAF1 (C57BL/6xCBA) male mice (aged 8-20 weeks) were used for the heterologous and homologous IVF experiments. The animals were maintained under conditions of 14 hours of light and 10 hours of darkness, with no food or water restrictions. Male mice were individually isolated for a period of 2 weeks before sperm collection. Euthanasia was performed by cervical dislocation. Each cauda epididymis was placed in M2 (M2 medium, Sigma-Aldrich M7167) to wash them out and remove the artery of the vas deferens, followed by washing out with HTF medium (human tubal fluid) (2.04 mM CaCl2, 101.6 mM NaCl, 4.69 mM KCl, 0.37 mM KH2PO4, 0.2 mM MgSO4, 21.4 mM sodium lactate, 0.33 mM sodium pyruvate, 2.78 mM glucose, 25 mM NaHCO3, 100 U/mL penicillin, 50 μg/mL streptomycin and 0.001% (w/v) phenol red); supplemented with 1% (w/v) BSA). Spermatozoa were collected by gently squeezing the dissected vas deferens and cauda epididymis, and were then incubated to swim-out in a 500 μL droplet of HTF medium, which supports capacitation, for 30 minutes at 37°C in the 5%CO2 incubator. The concentration of the sperm sample was determined using a Thoma cell counting chamber. Sperm motility quality was assessed following placement of 6 μL of sperm suspension in a Mackler® chamber on the stage of a microscope heated to 37°C (Nikon Eclipse E400). Fertilization was conducted in a final concentration of 2x10^5 capacitated spermatozoa per mL (56). All studies involving mice were approved by the Ethics Committee on Animal Experimentation of the INIA (Madrid, Spain), and were registered with the Dirección General de Agricultura y Ganadería of the Comunidad de Madrid (Spain) under PROEX 137.2/21.

Revised version: Sperm from B6CBAF1 (C57BL/6xCBA) male mice (aged 8-20 weeks) were used for the heterologous and homologous IVF experiments. The animals were maintained under conditions of 14 hours of light and 10 hours of darkness, with no food or water restrictions. Male mice were individually isolated for a period of 2 weeks before sperm collection. Euthanasia was performed by cervical dislocation. Each cauda epididymis was placed in M2 (M2 medium, Sigma-Aldrich M7167) to wash them out and remove the artery of the vas deferens, followed by washing out with HTF medium (human tubal fluid) (2.04 mM CaCl2, 101.6 mM NaCl, 4.69 mM KCl, 0.37 mM KH2PO4, 0.2 mM MgSO4, 21.4 mM sodium lactate, 0.33 mM sodium pyruvate, 2.78 mM glucose, 25 mM NaHCO3, 100 U/mL penicillin, 50 μg/mL streptomycin and 0.001% (w/v) phenol red); supplemented with 1% (w/v) BSA). Spermatozoa were collected by gently squeezing the dissected vas deferens and cauda epididymis, and were then incubated to swim-out in a 500 μL droplet of HTF medium, which supports capacitation, for 30 minutes at 37°C in the 5%CO2 incubator. The concentration of the sperm sample was determined using a Thoma cell counting chamber. Sperm motility quality was assessed following placement of 6 μL of sperm suspension in a Mackler® chamber on the stage of a microscope heated to 37°C (Nikon Eclipse E400). Fertilization was conducted in a final concentration of 2x10^5 capacitated spermatozoa per mL (60). All studies involving mice were approved by the Ethics Committee on Animal Experimentation of the INIA (Madrid, Spain), and were registered with the Dirección General de Agricultura y Ganadería of the Comunidad de Madrid (Spain) under PROEX 137.2/21.

Earlier version: Cat spermatozoa were collected from the cauda epididymis by modified retrograde flushing (57). Collected spermatozoa were incubated in 200 µL of medium, and their concentration determined using a Neubauer cell counting chamber. For frozen sperm, 50 µL of spermatozoa suspension (4 x 10^6 spermatozoa/mL) were stored in liquid nitrogen vapor for 2 hours to reach a temperature of 4°C, and then the samples were frozen and stored in liquid nitrogen at -196°C or lower. The frozen semen straws were thawed at 37°C in a water bath for 1 minute and then centrifuged for 8 minutes at 300 g using a gradient of 1 mL of 40% and 1 mL of 80% BoviPure (Nidacon Laboratories AB, Göthenborg, Sweden) according to the manufacturer’s instructions. The sperm pellet was isolated and washed in 3 mL of BoviWash (Nidacon Laboratories AB, Göthenborg, Sweden) by centrifugation at 300 g for 5 minutes. The pellet was re-suspended in the remaining 80 µL of BoviWash. Fertilization was carried out in a final concentration of 2x10^6 capacitated spermatozoa/mL.

Revised version: Cat spermatozoa were collected from the cauda epididymis by modified retrograde flushing following vasectomy (61). Collected spermatozoa were incubated in 200 µL of medium, and their concentration determined using a Neubauer cell counting chamber. For frozen sperm, 50 µL of spermatozoa suspension (4 x 10^6 spermatozoa/mL) were stored in liquid nitrogen vapor for 2 hours to reach a temperature of 4°C, and then the samples were frozen and stored in liquid nitrogen at -196°C or lower. The frozen semen straws were thawed at 37°C in a water bath for 1 minute and then centrifuged for 8 minutes at 300 g using a gradient of 1 mL of 40% and 1 mL of 80% BoviPure (Nidacon Laboratories AB, Göthenborg, Sweden) according to the manufacturer’s instructions. The sperm pellet was isolated and washed in 3 mL of BoviWash (Nidacon Laboratories AB, Göthenborg, Sweden) by centrifugation at 300 g for 5 minutes. The pellet was re-suspended in the remaining 80 µL of BoviWash. Fertilization was carried out in a final concentration of 2x10^6 capacitated spermatozoa/mL.

FITC-PNA Staining and visualization

The spermatozoa were diluted in PBS to a concentration of 2×10^6 spermatozoa/ml. Twenty μl of the mixture were placed on a clean glass slide and gently pressed between tissue paper sheets to remove excess liquid. The glass slides were heated to 37°C until the droplets of the sample dilution in PBS were dry. Once dry, the slides were immersed in cold methanol for 30 sec. and then left to dry for 15 min. at room temperature (RT). Next, two washes were performed in PBS for 5 min. To carry out the staining process, a standard dilution of 1:10 FITC-PNA (150 μg/ml) (L7381, Sigma-Aldrich, Madrid) in Hoechst (0.065 μg/ml) (B2883, Sigma-Aldrich) was prepared. Twenty μl of this solution were added to the spermatozoa sample on the glass slide. The stained samples were stored in a dark, humid box for a minimum of 30 min. After this time, the slides were washed in distilled water for 10 min, removing the excess and letting them dry on a heated plate at 37°C. Once dry, and keeping the slides in darkness, 2 μl of the commercial Fluoromount (F4680, Sigma-Aldrich) mixture were placed on the slide, covered with a coverslip (22×22 mm), and gently pressed between tissue paper sheets to remove excess liquid. After 30 min., the samples were ready for visualization. Spermatozoa were observed under a phase contrast microscope at 1,000× magnification with fluorescence illumination.

Homologous and heterologous IVF of bovine oocytes

Earlier version: Ovaries from mature heifers were collected at a local slaughterhouse, and immature cumulus–oocyte complexes (COCs) obtained by aspirating follicles of diameter 2 to 8 mm (58). Groups of 50 COCs were cultured in four-well dishes (Nunc) for 24 hours at 38.5°C under 5% CO2 in air, using 500 µL of maturation medium. This medium was composed of TCM 199 supplemented with 10% fetal calf serum (FCS) and 10 ng/mL of epidermal growth factor (EGF).

Revised version: Ovaries from mature heifers were collected at a local slaughterhouse, and immature cumulus–oocyte complexes (COCs) obtained by aspirating follicles of diameter 2 to 8 mm (62). Groups of 50 COCs were cultured in four-well dishes (Nunc) for 24 hours at 38.5°C under 5% CO2 in air, using 500 µL of maturation medium. This medium was composed of TCM 199 supplemented with 10% fetal calf serum (FCS) and 10 ng/mL of epidermal growth factor (EGF).

Earlier version: For homologous fertilization, frozen/thawed semen (0.25 mL) from a single Asturian Valley bull was treated by density gradient centrifugation (BoviPure®, Nidacon International, Sweden), and matured COCs where then co-incubated with selected sperm at a final concentration of 1 × 10^6 spermatozoa per mL for 18–22 hours in 500 μL of fertilization medium (59). This medium consisted of Tyrode’s medium containing 22 mM sodium lactate, 25 mM bicarbonate, 1 mM sodium pyruvate, and 6 mg/mL fatty acid-free BSA supplemented with 10 µg/mL heparin sodium salt (Calbiochem). Incubations took place in four-well dishes as groups of 50 COCs per well at 38.5°C under an atmosphere of 5% CO2 in air and maximum humidity.

Revised version: For homologous fertilization, frozen/thawed semen (0.25 mL) from a single Asturian Valley bull was treated by density gradient centrifugation (BoviPure®, Nidacon International, Sweden), and matured COCs where then co-incubated with selected sperm at a final concentration of 1 × 10^6 spermatozoa per mL for 18–22 hours in 500 μL of fertilization medium (63). This medium consisted of Tyrode’s medium containing 22 mM sodium lactate, 25 mM bicarbonate, 1 mM sodium pyruvate, and 6 mg/mL fatty acid-free BSA supplemented with 10 µg/mL heparin sodium salt (Calbiochem). Incubations took place in four-well dishes as groups of 50 COCs per well at 38.5°C under an atmosphere of 5% CO2 in air and maximum humidity.

For heterologous IVF, the bovine mature COCs were co-incubated with spermatozoa from other species (cat, human, and mouse). Sperm-egg interactions were assessed through a sperm-ZP binding assay at 2.5 hours post-incubation (hpi). For this purpose, COCs were vortexed for 3 minutes, fixed in 4% paraformaldehyde for 30 minutes, washed with cold PBS, and stained with Hoechst 33342 to count the number of sperm that remained bound to the ZP. This was done using a widefield fluorescence microscope with structured illumination (ApoTome, Zeiss) and UV-2E/C excitation: 340–380 nm and emission: 435– 485 nm. In addition, pronuclear formation was assessed at 6, 12, 18, and 22 hpi. For this, presumptive zygotes were treated and examined following the procedures explained above for sperm-ZP binding. The times used in this study for hybrid embryo analysis were similar to those used in other incubation studies performed in other species (18, 22).

In vitro culture of presumptive zygotes to first cleavage embryos on Day 2

After 22 hours incubation of bovine oocytes with sperm (bull, human, mice, or cat), presumptive zygotes or COCs were denuded of cumulus cells by vortexing for 3 minutes and then cultured in groups of approximately 20 in 25 µL droplets of synthetic oviductal fluid (SOF)(31). The SOF contained 4.2 mM sodium lactate, 0.73 mM sodium pyruvate, 30 mL/mL BME amino acids, 10 mL/mL minimum essential medium (MEM) amino acids, and 1 mg/mL phenol red. These cultures were supplemented with 5% FCS and placed under mineral oil in an atmosphere of 5% CO2, 5% O2, and 90% N2, and maximum humidity at 38.5°C (cow) or 37°C (human, mouse, and cat). Forty eight hours after the start of incubation, cleavage was assessed under a stereomicroscope by checking for the presence of two cells per embryo.

Collection of oviductal fluid

Five stage I (Days 1 to 4) oviducts and another five stage II (Days 5 to 9) were selected from slaughtered heifers. The stage of the estrous cycle was estimated based on the morphology of ovarian structures, especially the corpus luteum (CL). Briefly, in stage I, the CL is small and red, and the epithelium has not yet covered the point of follicular rupture. In stage II, the CL is somewhat bigger, vascularization is observed around it, the apex is reddish in color, and the epithelium has already covered the rupture point. Each tract was independently deposited in a plastic bag and transported to the laboratory on ice in a polystyrene box within 2 hours of collection. The oviducts were washed with cold water to remove blood traces and then in cold Ca2+ and Mg2+-free phosphate-buffered saline (PBS-). Tissue surrounding the oviduct was carefully isolated and the oviducts were flushed with 1 mL of cold PBS-from the ampulla to the isthmus. All manipulations were performed at 4°C. The fluid from the five oviducts was pooled and then centrifuged once at 300 × g for 7 minutes. Subsequently, the supernatant was centrifuged again at 10,000 × g for 30 minutes at 4°C to remove debris (32). The resulting supernatant was aliquoted and stored at -80°C.

Collection of oviductal fluid from mice

Estrous cycle stage was determined in mice by visual examination of the vaginal opening and cytological examination of the vagina smear. Oviducts from four estrous female mice were collected and washed by flushing from the ampulla to the isthmus with 0.2 mL of Ca2+ and Mg2+-free phosphate-buffered saline. Female mice were sacrificed sequentially and oviductal fluid was washed out with the same PBS- and oviducts kept on ice after each flushing. This was followed by storage at -80°C.

Preparation of empty zona pellucidae from bovine ovarian oocytes

COCs obtained from ovaries from slaughtered heifers and matured in vitro underwent a thorough washing process before their immersion in 30 mL of Tyrode’s medium. The oocytes were then mechanically denuded by vortexing and gentle pipetting. The denuded oocytes were subjected to three washing cycles using either G-IVF™ PLUS (Vitrolife, Igenomix) or FERT medium, depending on the experimental requirements. Batches of approximately 30 oocytes were then carefully placed in 400-μL droplets of the selected medium and covered with mineral oil. For the removal of cytoplasmic contents, a micromanipulator was used following the procedure outlined in Figure S4. The oocyte was secured using a glass capillary holder, while an ICSI needle, attached to a piezoelectric manipulator, was used to create a small opening in the ZP. This opening facilitated the penetration of the needle, allowing for the complete aspiration of the oocyte’s cytoplasm. The aspiration needle was then retrieved from the ZP to expel the extracted cytoplasm. Following this, the empty ZP underwent two washing cycles using G-IVF™ PLUS or FERT medium. These prepared zona pellucidae were either frozen at -20°C for a duration of several months (3-9 months) or directly utilized in the sperm penetration experiments.

Collection and in vitro maturation of ovarian cumulus-oocyte complexes in mice

Earlier version: Female mice aged 8 to 10 weeks were superovulated by intraperitoneal injections of 0.1 mL of PMSG (pregnant mare serum gonadotropin, Folligon, 5 IU, INIA, Madrid, Spain). Fifty hours after PMSG supplementation, the females were euthanized by cervical dislocation, and their ovaries removed. Collected ovaries were cleaned of any connective tissue and placed in M2 handling medium supplemented with 4 mg/mL of bovine serum albumin fraction V. Antral follicles were gently punctured with 30-gauge needles, and COCs with at least 3 compact cumulus cells layers were collected in handling medium and matured for 16-17 hours in TCM199 supplemented with 10% FCS and 10 ng/mL EGF (epidermal growth factor) (60) and kept at 37°C under an atmosphere of 5% CO2 in air with maximum humidity. Matured oocytes were isolated and placed in an appropriate medium; HTF for IVF and M2 for ZP isolation and later protein incubation.

Revised version: Female mice aged 8 to 10 weeks were superovulated by intraperitoneal injections of 0.1 mL of PMSG (pregnant mare serum gonadotropin, Folligon, 5 IU, INIA, Madrid, Spain). Fifty hours after PMSG supplementation, the females were euthanized by cervical dislocation, and their ovaries removed. Collected ovaries were cleaned of any connective tissue and placed in M2 handling medium supplemented with 4 mg/mL of bovine serum albumin fraction V. Antral follicles were gently punctured with 30-gauge needles, and COCs with at least 3 compact cumulus cells layers were collected in handling medium and matured for 16-17 hours in TCM199 supplemented with 10% FCS and 10 ng/mL EGF (epidermal growth factor) (64) and kept at 37°C under an atmosphere of 5% CO2 in air with maximum humidity. Matured oocytes were isolated and placed in an appropriate medium; HTF for IVF and M2 for ZP isolation and later protein incubation.

Isolation of superovulated mouse oocytes

Metaphase II oocytes were collected from the oviducts of 6-to 8-week-old female mice superovulated with 7.5 IU of PMSG, and an equivalent dose of hCG (human chorionic gonadotropin) 48 hours later. Briefly, 14 hours post-chorionic gonadotropin administration, oviducts were removed from the superovulated female mice and placed in a petri dish containing M2 at room temperature. After washing, collected oviducts were placed in fresh M2 medium, and COCs previously in contact with oviductal fluid were released from the ampulla with the aid of Dumont #55 forceps and washed in fresh M2 medium until cytoplasm removal.

Preparation of zona pellucidae from mouse oocytes

After obtaining the COCs using the two methods described, cumulus cells were dispersed by 3 to 5 minutes of incubation in M2 medium containing 350 IU/mL of hyaluronidase (Sigma-Aldrich H4272). After washing, oocytes were kept in KSOM medium at 37°C in an atmosphere of 5% CO2 until use.

To obtain empty ZP, everything except ZP (nucleus, organelles, membranes and cytoplasmic contents of the oocytes) was removed using a micromanipulator, following the procedure outlined in Figure S4. The protocol used was similar to that employed for the bovine oocytes with the help of a holding pipette and a blunt microinjection pipette. The decumulated oocytes were placed in a drop of M2 medium, and with the assistance of a piezoelectric manipulator, each zona pellucida was penetrated, followed by the removal of the cytoplasm using the blunt microinjection pipette.

Empty zona penetration test (EZPT) of murine and bovine zona pellucidas

For EZPTs, murine and bovine ZPs were firstly incubated with 30 ng/µL of recombinants OVGP1 (bovine, murine and human) for 30 minutes in 20µL drops. Then, ZP were washed out in media to remove the excess of recombinant OVGP1.

Earlier version: In some experiments the ZPs were exposed to oviductal fluid before the incubation with sperm from different species to analyze sperm penetration and sperm binding to the ZPs. In other experiments the ZPs were exposed or no exposed to one of the three recombinants OVGP1 (bovine, murine, and human), and each group were incubated with sperm of one of the three species using a final concentration of 1x10^6 sperm/mL. Male gametes were then incubated with ZPs for 4 hours in 50 µL of HTF medium for murine ZP or 50 µL of FERT medium for bovine ZPs, both covered with mineral oil in a 4-well plate at 38.5°C in a 5% CO2 maximum humidity atmosphere. EZPT was carried out under the same culture conditions for all groups. A positive control was set up through homologous fertilization with bovine or murine semen for each group.

Revised version: In some experiments the ZPs were exposed to oviductal fluid before the incubation with sperm from different species to analyze sperm penetration and sperm binding to the ZPs. In other experiments the ZPs were either exposed or not exposed to one of the three recombinants OVGP1 (bovine, murine, and human), and each group was incubated with sperm of one of the three species using a final concentration of 1x10^6 sperm/mL. Male gametes were then incubated with ZPs for 4 hours in 50 µL of HTF medium for murine ZP or 50 µL of FERT medium for bovine ZPs, both covered with mineral oil in a 4-well plate at 38.5°C in a 5% CO2 maximum humidity atmosphere. EZPT was carried out under the same culture conditions for all groups. A positive control was set up through homologous fertilization with bovine or murine semen for each group.

Also, in some experiments, bovine and murine capacitated sperm samples were incubated with OVGP1 under similar conditions to those used for ZPs (incubated with 30 ng/µL of recombinant OVGP1 (bovine and murine) for 30 minutes in 50 µL drops). To avoid contamination by OVGP1 present in the medium of the sperm samples, these were diluted in the corresponding medium (FERT medium for bovine, HTF for mice) and centrifuged at 1,800 x g for 10 minutes. The supernatant was discarded, and the samples were resuspended in the same media. This procedure was performed twice for each sample. After that, the sperm concentration was adjusted and used as previously described in the protocol for EZPT. Four hours after incubation, the ZPs were transferred from the 50 µL drop of fertilization medium to a new drop to correctly assess both sperm binding and penetration to the ZP. Once the zonas were in the new drop, the spermatozoa binding or inside every zona pellucida were quantified. In each incubation, at least 30 ZPs were used per group and three repetitions were performed for each experimental group. The exact number of ZPs used per group can be seen in the supplementary Tables S6 and S8-13.

Origins of bovine, murine and human OVGP1 recombinants

Earlier version: The coding sequences of bovine OVGP1 (NM_001080216.1), an additional leader peptide sequence for mammalian cell expression (MGWSCIILFLVATATGVHS) and the sequence to express six-His residues were cloned into a pcDNA3.1(+) expression vector from GenScript Giotech (Rijswijk, NL) containing a promoter for bacteriophage T7 RNA polymerase. Tagged OVGP1 protein was produced by transient expression using a Vaccinia/T7 system (61) in BHK-21 cells (ATCC CCL-10) grown in G-MEM BHK-21 medium (Gibco) supplemented with 5% FBS (Gibco), 10% tryptose phosphate medium (Gibco), 20 mM Hepes (Fisher Scientific), and supplemented with 2 mM L-glutamine (Lonza), 0.1 µg/ml penicillin, and 0.1 µg/ml streptomycin (Lonza). BHK-21 cells grown in 6-well plates were infected with a crude stock of vaccinia virus expressing T7 polymerase (vTF7-3) (62) at a multiplicity of infection (m.o.i.) of 5 pfu (plaque-forming unit)/cell in 2% FBS EMEM (Eagle’s minimum essential medium) medium. After a 90-minute adsorption period, the virus inoculum was removed, cell monolayers were washed with EMEM medium and subsequently transfected with pH224 plasmid. Transfection for each well was performed with a mixture of 2 µg of DNA and 6 µL of Fugene HD reagent (Promega). The DNA/fugene mixture was incubated for 15 minutes at room temperature prior to its drop-wise addition to the infected culture. At 3 days post-infection, the cell culture was harvested and cells were sedimented by low speed centrifugation.

Revised version: The coding sequences of bovine OVGP1 (NM_001080216.1), an additional leader peptide sequence for mammalian cell expression (MGWSCIILFLVATATGVHS) and the sequence to express six-His residues were cloned into a pcDNA3.1(+) expression vector from GenScript Giotech (Rijswijk, NL) containing a promoter for bacteriophage T7 RNA polymerase. Tagged OVGP1 protein was produced by transient expression using a Vaccinia/T7 system (65) in BHK-21 cells (ATCC CCL-10) grown in G-MEM BHK-21 medium (Gibco) supplemented with 5% FBS (Gibco), 10% tryptose phosphate medium (Gibco), 20 mM Hepes (Fisher Scientific), and supplemented with 2 mM L-glutamine (Lonza), 0.1 µg/ml penicillin, and 0.1 µg/ml streptomycin (Lonza). BHK-21 cells grown in 6-well plates were infected with a crude stock of vaccinia virus expressing T7 polymerase (vTF7-3) (66) at a multiplicity of infection (m.o.i.) of 5 pfu (plaque-forming unit)/cell in 2% FBS EMEM (Eagle’s minimum essential medium) medium. After a 90-minute adsorption period, the virus inoculum was removed, cell monolayers were washed with EMEM medium and subsequently transfected with pH224 plasmid. Transfection for each well was performed with a mixture of 2 µg of DNA and 6 µL of Fugene HD reagent (Promega). The DNA/fugene mixture was incubated for 15 minutes at room temperature prior to its drop-wise addition to the infected culture. At 3 days post-infection, the cell culture was harvested and cells were sedimented by low speed centrifugation.

Bovine OVGP1 (bOVGP1) recombinant protein was purified using the Amicon Pro Purification System (Merck Millipore Ltd) by combining metal chelation chromatography with an Amicon concentrator membrane. Mammalian cell supernatants were incubated under constant agitation for 2 hours with the Ni-NTA His-Bind resin, previously equilibrated with binding buffer. Then, following the manufacturer’s protocol specifications and after washing and elution centrifugation steps, His-tagged proteins were concentrated in an Amicon Ultra-0.5 device, and eluted in 25 μL of the desired buffer (50 mM Tris-HCl, 500 mM NaCl, 5% glycerol, pH=8) and stored at -80°C.

Human (NM_002557.4) and mouse (NM_007696.2) recombinant oviductin proteins were obtained from Origene Technologies Inc. (human Cat. No. TP321684, murine Cat. No. Tp521693; Rockville, US). A sequence expressing FLAG-tagged epitope proteins (DYKDDDDK) was cloned into an expression vector. The recombinant proteins were produced by transient transfection in HEK293T, obtained from the cell culture supernatant, and captured on an affinity column followed by conventional chromatography steps.

Western blotting of OVGP1

Earlier version: Proteins were run on an SDS-PAGE gel of 10% acrylamide, loading per well 300 -500 ng of total purified commercial protein from Origene, an equivalent amount of the in-house produced proteins from oviductal fluid, and 10 μL of bovine and 50 μL of murine proteins. . Protein ladders used were the Broad Multi Color Pre-Stained Protein Standard (M00624S, GenScript) and All Blue prestained Precision Plus Protein Standard (Bio-Rad) The resultant gel was transferred to a nitrocellulose membrane for immunoblotting following standard procedures. Protein loading was assessed after blotting by Ponceau S Solution (Sigma-Aldrich,P7170). Membranes were blocked with 3% BSA in PBS-T (0.05% Tween in 1X PBS) and incubated with shaking overnight at 4°C with anti-OVGP1 rabbit monoclonal antibody (NBP1-76939) (1:750) for OVGP1, and anti-Histidine Tag antibody rabbit monoclonal clone RM146 (SAB5600227, Sigma-Aldrich) (1:1000) for recombinant bOVGP1 and mOVGP1, and Anti-Flag M2 antibody mouse monoclonal 639 (F1804, Sigma-Aldrich) (1:1000) for recombinant hOVGP1. Incubation was continued for 2 hours at room temperature using as secondary antibodies goat anti-rabbit IgG-HRP (Cat. No. GTX213110-01, GeneTex), or goat anti-mouse mIgGk-HRP (sc.516102, Santa Cruz), both diluted 1:5000. As loading control, was used the primary antibody monoclonal anti-vinculin rabbit antibody (926-42215, LICORbio). Membranes were developed with Immobilon® Forte Western HRP Substrate (Cat. No. WBLUF0100, Millipore). The chemiluminescence signal was digitalized using an ImageQuant LAS 500 chemiluminescence CCD camera (GE Healthcare Life Sciences, USA, 29005063).

Revised version: Proteins were run on an SDS-PAGE gel of 10% acrylamide, loading per well 300 -500 ng of total purified commercial protein from Origene, and an equivalent amount of the bovine in-house produced protein. From oviductal fluid were used 10 μL of bovine and 50 μL of murine proteins. Protein ladders used were the Broad Multi Color Pre-Stained Protein Standard (M00624S, GenScript) and All Blue prestained Precision Plus Protein Standard (Bio-Rad) The resultant gel was transferred to a nitrocellulose membrane for immunoblotting following standard procedures. Protein loading was assessed after blotting by Ponceau S Solution (Sigma-Aldrich,P7170). Membranes were blocked with 3% BSA in PBS-T (0.05% Tween in 1X PBS) and incubated with shaking overnight at 4°C with anti-OVGP1 rabbit monoclonal antibody (NBP1-76939) (1:750) for OVGP1 (able to bind to recombinant human, murine and bovine OVGP1 and OVGP1 from oviductal fluids), with anti-Histidine Tag antibody rabbit monoclonal clone RM146 (SAB5600227, Sigma-Aldrich) (1:1000) for recombinant bOVGP1, and with Anti-Flag M2 antibody mouse monoclonal 639 (F1804, Sigma-Aldrich; St. Louis, MO, USA) (1:1000) for recombinant hOVGP1 and mOVGP1. Incubation was continued for 2 hours at room temperature using as secondary antibodies goat anti-rabbit IgG-HRP (Cat. No. GTX213110-01, GeneTex), or goat anti-mouse mIgGk-HRP (sc.516102, Santa Cruz), both diluted 1:5000. As loading control, the primary antibody monoclonal anti-vinculin rabbit antibody (926-42215, LICORbio) was used. Membranes were developed with Immobilon® Forte Western HRP Substrate (Cat. No. WBLUF0100, Millipore). The chemiluminescence signal was digitalized using an ImageQuant LAS 500 chemiluminescence CCD camera (GE Healthcare Life Sciences, USA, 29005063).

The total protein amount from oviductal fluids was determined using the PierceTM Protein Assay Kit. The concentration of OVGP1 in the oviductal fluid was quantified using ImageJ software by comparing the mean gray value of the band in the oviductal fluid to the band in the recombinant protein lane. By establishing this relationship, along with the known concentration of protein amount in the recombinant one and in the total protein amount of oviductal fluid, the concentration of OVGP1 in the oviductal fluid was determined as the average of three western blots.

OVGP1 immunofluorescence

Earlier version: Zona pellucidae were incubated with OVGP1 proteins from human, bovine, and murine (at final concentration of 20ng/µL), or with murine oviductal fluid for 30 minutes at room temperature. Then, they were fixed in 4% PFA (paraformaldehyde) in PBS 1X 1% BSA for 20 minutes at room temperature. The zonas were washed out twice with PBS 1X 1% BSA and incubated with blocking buffer (5% FBS in PBS 1X) for 45 minutes. After blocking, primary antibody was prepared at 1:50 in 20% blocking buffer in PBS 1X 1% BSA followed by incubation in a humid atmosphere at 4°C overnight. ZPs were washed out twice in PBS 1X 1% BSA and incubated for 2 hours at room temperature with the secondary antibody 1:200 (in 20% blocking buffer in PBS 1X 1% BSA). After 3 wash-outs in PBS 1X 1% BSA, the ZPs were embedded in Fluoromount Aqueous Mounting Medium (F4680, Sigma-Aldrich) and observed under structured illumination in a Zeiss Axio Observer microscope equipped with ApoTome.2 software. The primary antibodies used were anti-OVGP1, rabbit monoclonal (NBP1-76939) for endogenous OVGP1 and recombinant bOVGP1 and mOVGP1 and Anti-Flag M2 antibody, mouse monoclonal (F1804, Sigma-Aldrich) for recombinant hOVGP1. The secondary antibodies used were Alexa Fluor 647 goat anti-rabbit (H+L) (A21245, Invitrogen) and Alexa Fluor 488 goat anti-mouse IgG (H+L) (A11029, Invitrogen) (63).

Revised version: Zona pellucidae were incubated with OVGP1 proteins from human, bovine, and murine (at final concentration of 20ng/µL), or with murine oviductal fluid for 30 minutes at room temperature. Then, they were fixed in 4% PFA (paraformaldehyde) in PBS 1X 1% BSA for 20 minutes at room temperature. The zonas were washed out twice with PBS 1X 1% BSA and incubated with blocking buffer (5% FBS in PBS 1X) for 45 minutes. After blocking, primary antibody was prepared at 1:50 in 20% blocking buffer in PBS 1X 1% BSA followed by incubation in a humid atmosphere at 4°C overnight. ZPs were washed out twice in PBS 1X 1% BSA and incubated for 2 hours at room temperature with the secondary antibody 1:200 (in 20% blocking buffer in PBS 1X 1% BSA). After 3 wash-outs in PBS 1X 1% BSA, the ZPs were embedded in Fluoromount Aqueous Mounting Medium (F4680, Sigma-Aldrich) and observed under structured illumination in a Zeiss Axio Observer microscope equipped with ApoTome.2 software. The primary antibodies used were anti-OVGP1, rabbit monoclonal (NBP1-76939) for endogenous OVGP1 and recombinant bOVGP1 and mOVGP1 and Anti-Flag M2 antibody, mouse monoclonal (F1804, Sigma-Aldrich) for recombinant hOVGP1. The secondary antibodies used were Alexa Fluor 647 goat anti-rabbit (H+L) (A21245, Invitrogen) and Alexa Fluor 488 goat anti-mouse IgG (H+L) (A11029, Invitrogen) (67).

Proteomics identification of murine OVGP1

Earlier version: Oviductal fluid from 5 female mice was obtained in PBS as previously described, and 50 µL aliquots loaded on an 8% polyacrylamide gel for SDS-PAGE. Some lanes of the gel were blotted onto a nitrocellulose membrane for immunoblotting against OVGP1 antibody, and other lanes were fixed for 30 minutes in methanol:acetic acid (50:10) and rinsed with distilled water. Fixed bands were stained with Coomassie Brilliant Blue, and 75kDa and 50kDa bands (compared to the lanes immunoblotted for OVGP1) were cut, destained in acetonitrile:water (ACN:H2O, 1:1), digested with trypsin, and subjected to reverse phase-liquid chromatography RP-LC-MS/MS mass scan analysis (64). Murine oviductin identification was performed with the PEAKS Studio v11.5 search engine (Bioinformatics Solutions Inc., Waterloo, Ontario, Canada) against the Uniprot Mus musculus database using a filter of peptides and proteins with a FDR<=1% (65).

Revised version: Oviductal fluid from 5 female mice was obtained in PBS as previously described, and 50 µL aliquots loaded on an 8% polyacrylamide gel for SDS-PAGE. Some lanes of the gel were blotted onto a nitrocellulose membrane for immunoblotting against OVGP1 antibody, and other lanes were fixed for 30 minutes in methanol:acetic acid (50:10) and rinsed with distilled water. Fixed bands were stained with Coomassie Brilliant Blue, and 75kDa and 50kDa bands (compared to the lanes immunoblotted for OVGP1) were cut, destained in acetonitrile:water (ACN:H2O, 1:1), digested with trypsin, and subjected to reverse phase-liquid chromatography RP-LC-MS/MS mass scan analysis (68). Murine oviductin identification was performed with the PEAKS Studio v11.5 search engine (Bioinformatics Solutions Inc., Waterloo, Ontario, Canada) against the Uniprot Mus musculus database using a filter of peptides and proteins with a FDR<=1% (69).

Scanning Electron Microscopy and Image Processing

Earlier version: Oocytes were denuded as detailed above and prepared for SEM as previously described (66) with minor modifications. Briefly, oocytes were placed in fixation medium (2.5% glutaraldehyde [v/v] and 0.1 mol/L sodium cacodylate buffer) for 1 hour at 4°C. They were then washed with 0.1 mol/L sodium cacodylate buffer and kept in the buffer for 1 hour at 4°C, followed by washing in distilled water for 5 minutes. Thereafter, the oocytes were dehydrated with increasing concentrations of acetone at room temperature. After their dehydration, they were CO2-critical point dried in a stainless chamber (Denton Vacuum DCP-1, Denton Vacuum, LLC, Moorestown, NJ, USA), mounted onto aluminum stubs, and coated with platinum sputter-coated with 5.0 nm thin layer (Leica EM ACE 600). Photomicrographs of the ZP of oocytes were taken a FESEM electron microscope. FESEM analyses were performed with a FE-SEM (ApreoS Lovac IML Thermofisher).

Revised version: Oocytes were denuded as detailed above and prepared for SEM as previously described (70) with minor modifications. Once denuded, oocytes were either exposed or not exposed to one of the three recombinants OVGP1. Briefly, oocytes were placed in fixation medium (2.5% glutaraldehyde [v/v] and 0.1 mol/L sodium cacodylate buffer) for 1 hour at 4°C. They were then washed with 0.1 mol/L sodium cacodylate buffer and kept in the buffer for 1 hour at 4°C, followed by washing in distilled water for 5 minutes. Thereafter, the oocytes were dehydrated with increasing concentrations of acetone at room temperature. After their dehydration, they were CO2-critical point dried in a stainless chamber (Denton Vacuum DCP-1, Denton Vacuum, LLC, Moorestown, NJ, USA), mounted onto aluminum stubs, and coated with platinum sputter-coated with 5.0 nm thin layer (Leica EM ACE 600). Photomicrographs of the ZP of oocytes were taken a FESEM electron microscope. FESEM analyses were performed with a FE-SEM (ApreoS Lovac IML Thermofisher).

Incubation of ZP, OVGP1, and sperm with neuraminidase (NMase)

Earlier version: Two hundred isolated ZPs were treated with Clostridium perfringens neuraminidase (type V, Sigma-Aldrich, N2876) in 0.1M acetate buffer at pH 4.5, as three separate replicates. The enzyme was tested at concentrations of 5 U/mL for an incubation time of 17 hours at 38.5°C.

Revised version: Two hundred isolated bovine ZPs and another 200 isolated murine ZPs were treated with Clostridium perfringens neuraminidase (type V, Sigma-Aldrich, N2876) in 0.1M acetate buffer at pH 4.5, as three separate replicates. The enzyme was tested at concentrations of 5 U/mL for an incubation time of 17 hours at 38.5°C. As control, ZPs were incubated under the same conditions in the absence of NMase. The NMase used shows broad specific activity, cleaving terminal sialic acid residues which are α-2,3-α-2,6-or α-2,8-linked to Gal, GlcNac or GalNAc oligosaccharides, glycolipids or glycoproteins (71).

Removed in the revised version: As control, ZPs were incubated under the same conditions in the absence of NMase. The NMase used shows broad specific activity, cleaving terminal sialic acid residues which are α-2,3-α-2,6-or α-2,8-linked to Gal, GlcNac or GalNAc oligosaccharides, glycolipids or glycoproteins (67).

When bovine and murine sperm were incubated with recombinant OVGP1, the same conditions used for ZP incubation were applied. To prevent OVGP1 from sperm incubation from being included in subsequent IVF procedures, semen samples were washed in Semen Wash Medium (Stroebech Media, Denmark) by discarding the supernatant after 5 minutes of centrifugation. After washing, the samples were prepared for EZPT with a final concentration of 1 × 10^6 sperm/mL.

Earlier version: When bovine OVGP1 protein alone was incubated with NMase, the same conditions used for ZP incubation were applied. OVGP1 was treated with Clostridium perfringens neuraminidase (Type V, Sigma) in 0.1M acetate buffer at pH 4.5, with three independent replicates for each experiment. The enzyme was tested at a concentration of 5 U/mL with an incubation time of 17 hours at 38.5°C. As a control, OVGP1 was incubated under the same conditions but without NMase. The NMase used exhibits broad specific activity, cleaving terminal sialic acid residues that are α-2,3-, α-2,6-, or α-2,8-linked to Gal, GlcNAc, or GalNAc oligosaccharides, glycolipids, or glycoproteins (67). After incubation of OVGP1 with the enzyme, OVGP1 was isolated from the NMase by incubation under constant agitation for 3 hours with the Ni-NTA His-Bind resin, followed by chromatography and concentration through an Amicon membrane, and then treated bOVGP1 was eluted in 25 L of the desired buffer (50 mM Tris-HCl, 500 mM NaCl, 5% glycerol, pH=8). After elution, ZPs were normally incubated for 30 min with treated or control bovine OVGP1 to perform EZPT.

Revised version: When bovine OVGP1 protein alone was incubated with NMase, the same conditions used for ZP incubation were applied. OVGP1 was treated with Clostridium perfringens neuraminidase (Type V, Sigma) in 0.1M acetate buffer at pH 4.5, with three independent replicates for each experiment. The enzyme was tested at a concentration of 5 U/mL with an incubation time of 17 hours at 38.5°C. As a control, OVGP1 was incubated under the same conditions but without NMase. The NMase used exhibits broad specific activity, cleaving terminal sialic acid residues that are α-2,3-, α-2,6-, or α-2,8-linked to Gal, GlcNAc, or GalNAc oligosaccharides, glycolipids, or glycoproteins (71). After incubation of OVGP1 with the enzyme, OVGP1 was isolated from the NMase by incubation under constant agitation for 3 hours with the Ni-NTA His-Bind resin, followed by chromatography and concentration through an Amicon membrane, and then treated bOVGP1 was eluted in 25 L of the desired buffer (50 mM Tris-HCl, 500 mM NaCl, 5% glycerol, pH=8). After elution, ZPs were normally incubated for 30 min with treated or control bovine OVGP1 to perform EZPT.

Statistical analysis

Statistical analysis was carried out using the software package SigmaStat (Jandel Scientific, San Rafael, CA). Results are provided as means ± standard deviation (SD). Means were compared and analyzed by repeated measures one-way analysis of variance (ANOVA), followed by Tukey’s post hoc test. Significance was set at p < 0.05.